<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">

# 第二回レポート課題の準備

</div>

このシートでは以下を整備します。

1. Q学習エージェント: `class Qlearner` (作成済みのものをコピペで持ってくる)
2. 強化学習の実行用関数: `play_env()`(作成済みのものをコピペで持ってきて，少し修正)
3. 結果の可視化: `plot_history()`(作成済みのものをコピペで持ってくる)
4. 結果の分析: `mean_performance()`(`play_env()`を何度も実行して，結果の平均と標準偏差を求める)
5. 結果の可視化その2: `plot_history()`(コピペで持ってきたのを少し変えて，標準偏差も表示できるようにする)
6. SARSA学習エージェント `class SARSAlearner` (すでに作ったものをコピペで持ってくる)
7. Actor-Critic学習エージェント(`class Qlearner`を修正して作る)
8. `MDPGridworld-v3`で遊ぶ

まずは，これまでと同様に`MDPGridworld-v1`環境を用います。

```python
# このセルでは，実行に必要なモジュールの読み込みを行う。
# ここに書かれているもの以外に必要なものがあれば，適宜追加すること。
import numpy as np
import random as ra
import matplotlib.pyplot as plt

import gym
```

<div class="alert alert-success">

# Q学習エージェント

</div>

**Q.** 以前の課題シート(gridworld)で作ったクラス`Qlearner`を入力しなさい。独自仕様のインスタンス関数等が追加されていても良いが，行動選択にはε-greedy法を使っているものを入力すること。

<!-- grade_id: cell-1485191ebed626ef, points: 0, grade: false, solution: true -->
```python
# 以下にclass QLearnerの定義を書きなさい。(必要ならindex_list()等, Qlearnerの実行に必要な関数も書きなさい)
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-success">

# 強化学習の実行用関数

</div>

## エージェントを動かす関数: `play_env()`
**Q.** 以前の課題シート(gridworld)で作った関数`play_env()`を，以下のように修正して，次のセルに入力しなさい。

**変更点**
- 変数`history_steps`, `history_score`にはリストではなく，**NumPyの`ndarray`を使う**。(全要素を0に初期化するための命令を`np.zeros()`を使って書き直す)

<!-- grade_id: cell-c910cfdb8bbb87ce, points: 0, grade: false, solution: true -->
```python
def play_env(env, agent, learning=True, max_episode=30, render_step=20, verbose=False):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# このセルで試してみよう
```

<!-- grade_id: E-check-play_env, points: 2, grade: true, solution: false -->
```python
# 評価用セル(返り値がndarrayか確認)
env = gym.make("MDPGridworld-v1")
action_list = list(range(env.action_space.n))  # action_list取得
agent = Qlearner(action_list)  # 学習エージェント生成

step, score = play_env(env=env, agent=agent, max_episode=100, verbose=False)
eq_(type(step), type(np.ones(1)))
eq_(type(score), type(np.ones(1)))
```

## 結果の可視化: `plot_history()`

Q. 次のセルにgridworld課題シートで作った関数`plot_history()`を入力して動作確認を行い，必要に応じて修正をしなさい。

<!-- grade_id: cell-4937f30c0d392b0b, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt


# YOUR CODE HERE
raise NotImplementedError()
```

Q. 次のセルを実行して，ここまでに入力した内容の動作確認をしなさい。

```python
# このセルで動作確認をしよう
# Qlearner + play_env() + MGPGridWorld の学習プログラム

env = gym.make("MDPGridworld-v1")
action_list = list(range(env.action_space.n))  # action_list取得
agent = Qlearner(action_list)  # 学習エージェント生成

step, score = play_env(env=env, agent=agent, max_episode=100, verbose=False)
plot_history(step, score)
```

<!-- grade_id: cell-4c6632fdd17d4004, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## 結果の分析: `mean_performance()`

Q. 上記のプログラムを何度か実行するとわかるように，強化学習には乱数を用いるので，実行するたびに若干結果が異なる。そこで，学習性能の評価のためには，関数`play_env()`で実行した結果を何度も実行し，返り値の`step`, `score`の平均値と標準偏差を評価する必要がある。このような評価を行うための関数`mean_performance()`を以下のように作成しなさい。


**引数とデフォルト値**

- `env`: 学習環境。`gym.make("MDPGridworld-v1")`の返り値などを渡す
- `agent`: 環境を動く行動を決定するエージェント。`Qlearner`などのインスタンス変数を渡す
- `n_trial`: `play_env()`を実行する回数(デフォルト値`10`)
- `max_episode`: 繰り返すエピソード数 (デフォルト値`100`)

**処理**

- `play_env()`を引数で指定した回数(`n_trial`)だけ実行する。`play_env()`の引数には，`mean_performance()`の引数(`env`, `agent`, `max_episode`)をそのまま渡す他，`learning=True`, `verbose=False`を指定する。
- **新たに`play_env()`を呼び出す毎に`agent.reset()`によって学習内容を消去**すること。
- 上記の実行の結果，`play_env()`の返り値`history_step`, `history_score`が`n_trial`回得られる。各データにおける第1,2,...,`max_episode`回目のエピソードのそれぞれについて平均値と標準偏差を求める。
- 返り値平均値と標準偏差は長さ`max_episode`の`ndarray`になることに注意
- **ヒント**) `numpy`課題シートにある`sinrnd()`の平均値を求める問を参考にすると良い

        
**返り値**

- `step_mean`: 各エピソードにおいてエピソード終了までにかかったステップ数の平均値
- `score_mean`: 各エピソードにおいて獲得した総報酬の平均値
- `step_sd`:  各エピソードにおいてエピソード終了までにかかったステップ数の標準偏差
- `score_sd`: 各エピソードにおいて獲得した総報酬の標準偏差

**実行例**
```python
agent = Qlearner(action_list)
step_mean, score_mean, step_sd, score_sd = mean_performance(
    env=env, agent=agent, n_trial=10, max_episode=100
)
```

<!-- grade_id: cell-e6ced5234d014b2a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-mean_performance, points: 4, grade: true, solution: false -->
```python
# 評価用セルその2
```

## 結果の可視化その2: `plot_history()`の修正

Q. 関数`plot_history()`を以下のように修正して、`mean_performance()`で得られた`steps`と`scores`の`mean±SD`のグラフを表示できるようにしなさい。

**引数**

- `mean_performance()`の4つの返り値を引数として渡せるように、引数には`step`, `score`に以下を追加。
    - `step_sd` : デフォルト値は`None` 
    - `score_sd`: デフォルト値は`None`
    - 注意: 
        - ある変数値が`None`である場合，その**変数が未定義**であることを表す。なお，`None`は**文字列ではない**ことに注意。
        - `if`文で「変数`a`が未定義でない」場合の条件分岐を書くには以下のように書く。
        ```python
        if a is not None:
            ...
        ```

**処理**

- 元の`plot_history()`と同様に引数`step`と`score`に対するグラフをそれぞれ座標コンテナ`ax1`と`ax2`に登録して表示する他、以下の例のように、`step±step_sd`と`score±score_sd`のグラフもそれぞれの上下に表示する。
- ただし，`step_sd`や`score_sd`の値が`None`の場合には，該当する標準偏差のグラフは表示しない。

**返り値**

- 元の`plot_history()`と同様に座標コンテナ`ax1`と`ax2`を返り値とする。


**実行例**
```python
agent = Qlearner(action_list)
step_mean, score_mean, step_sd, score_sd = mean_performance(env=env, agent=agent, n_trial=10, max_episode=100)
plot_history(step_mean, score_mean, step_sd, score_sd)
```

![](images/plot_history2.png)

<!-- grade_id: cell-e0f8d259f7107d2e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで試してみよう
agent = Qlearner(action_list)
step_mean, score_mean, step_sd, score_sd = mean_performance(env=env, agent=agent, n_trial=10, max_episode=100)
plot_history(step_mean, score_mean, step_sd, score_sd)
```

<!-- grade_id: E-plot_history2, points: 3, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# SARSA学習エージェント

</div>

Q. 課題シート'gridworld'で，`class Qlearner`の継承クラスとして作成した`class SARSAlearner`を次のセルにコピーしなさい。

<!-- grade_id: cell-978e927d584acc4f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
agent = SARSAlearner(action_list)
step_mean, score_mean, step_sd, score_sd = mean_performance(env=env, agent=agent, n_trial=10, max_episode=100)
plot_history(step_mean, score_mean, step_sd, score_sd)
```

<!-- grade_id: E-SARSA, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# Actor-Criticエージェント

</div>

**Q.** `class Qlearner`を改造して、Actor-Critic法による学習を行う`class ActorCritic`を作りなさい。**変更点は以下**になる。
`class Qlearner`の継承クラスとして作成しても良いが，新規に作り直しても良い。


- 初期化メソッドの修正・追加
    - 引数:
        - `v0`を追加。状態価値関数$V(s)$の初期値, デフォルト値は0.0
        - `T`を追加。行動選択に用いるsoftmax法のパラメータ, デフォルト値は1.0
        - 行動選択にはsoftmax法を利用するので，ε-greedy法のパラメータに用いていた引数`epsilon`は不要になる。ActorCriticの初期化メソッドの引数からは削除。
    - 処理内容
        - インスタンス変数`v0`に引数`v0`の値を代入
        - 状態価値関数$V(s)$の値を格納する空の辞書`V`を準備。(Qlearnerの辞書`Q`の初期化と同様)
        - その他の処理は`Qlearner`と同じ
- インスタンスメソッド`reset()`の追加処理
    - 辞書`V`の要素を空にリセットする(辞書`Q`のリセットと同様)
- インスタンスメソッド`get_V()`の追加
    - 引数: `s` (状態)
    - 返値: 辞書`V`を参照し、価値関数$V(s)$の値を返す。引数で与えられた状態`s`が辞書`V`のキーにないときには、インスタンス変数`v0`の値を返す。(すでにある`get_Q()`とほぼ同様の内容)
- インスタンスメソッド`learning()`の修正
    - Q学習をActor-Critic法に変更する。
        $$\delta=r+\gamma V(s')-V(s)$$
        $$V(s) \leftarrow V(s)+\alpha\delta$$
        $$Q(s, a) \leftarrow Q(s,a) + \alpha\delta$$
        ただし，`s'`がエピソード終了状態ならば
        $$\delta=r-V(s)$$
    - $Q(s,a)$は、状態$(s,a)$における行動を決定するための関数
        - 講義では関数`p(s,a)`と説明していた関数。Q学習の状態行動価値関数とは定義が異なるが，`Qlearner`で使っていた辞書`Q`をそのまま用いることにする。
    - $Q(s,a)$の学習定数$\alpha$は、$V(s)$の学習定数とは異なる値にすることもあるが、とりあえず同じ値にしておく。
- 行動選択メソッド`get_action()`の修正
    - $Q(s,a)$に基づく行動選択をsoftmax法にする。すなわち，行動`a`の選択確率は次式で与えられる。
        $$\displaystyle P(a)=\frac{e^\frac{Q(a)}{T}}{\displaystyle\sum_{i=1}^{N}e^\frac{Q(i)}{T}}$$
    - softmax法による行動決定に使う関数`softmax_rand()`には，課題シート`numpy`ですでに作成したものを用いなさい。

<!-- grade_id: cell-d728b712c482564c, points: 0, grade: false, solution: true -->
```python
# このセルに課題シート`numpy`で作成した関数softmax_rand()を入力する。
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9777873e0e4cd7ee, points: 0, grade: false, solution: true -->
```python
# このセルにclass ActorCriticの定義を入力する

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5953d114bf063069, points: 0, grade: false, solution: true -->
```python
# このセルで，`play_env()`と`plot_history()`を用いて動作確認をしよう
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-ActorCritic1, points: 1, grade: true, solution: false -->
```python
# 評価用セル1
action_list = ["left", "right"]
agent = ActorCritic(action_list,
        alpha=0.2,
        gamma=0.9,
        v0=0.0,
        q0=0.0,
        T=1.0)

# check init()
eq_(agent.alpha, 0.2)
eq_(agent.gamma, 0.9)
eq_(agent.q0, 0.0)
eq_(agent.v0, 0.0)
eq_(agent.T, 1.0)

eq_(agent.Q, {})
eq_(agent.V, {})
```

<!-- grade_id: E-ActorCritic2, points: 1, grade: true, solution: false -->
```python
# 評価用セル2
# check learning()
agent.reset()
agent.learning(s1=4, a1="right", r=-1, s2=5)
agent.learning(s1=5, a1="right", r=10, s2=6)
agent.learning(s1=3, a1="right", r=-1, s2=4)
agent.learning(s1=4, a1="right", r=-1, s2=5)
agent.learning(s1=5, a1="left", r=-1, s2=4)
agent.learning(s1=4, a1="right", r=-1, s2=5)
agent.learning(s1=5, a1="right", r=10, s2=6)
#print(agent.Q)
#print(agent.V)

ans_Q={(4, 'right'): 0.052000000000000005, (5, 'right'): 3.7199999999999998, (3, 'right'): -0.236, (5, 'left'): -0.6000000000000001}
ans_V={4: 0.052000000000000005, 5: 3.12, 3: -0.236}

assert all( (k,v) in agent.Q.items()
            for (k,v) in ans_Q.items()
)
assert all( (k,v) in agent.V.items()
            for (k,v) in ans_V.items()
)
```

<!-- grade_id: E-ActorCritic3, points: 1, grade: true, solution: false -->
```python
# 評価用セル
# check get_action()
```

<!-- grade_id: E-ActorCritic4, points: 1, grade: true, solution: false -->
```python
# 評価用セル
# check reset()
agent.reset()
eq_(agent.Q, {})
eq_(agent.V, {})
```

<div class="alert alert-success">

# MDPGridworld-v3

</div>

これまでのMDPGridworld-v1に，プチご褒美がおいてあるのがv3です。

```python
import gym

env = gym.make("MDPGridworld-v3")
env.show_info()
```

上の図のように，`C`と書かれているケーキがあります。
ここに来ると美味しいケーキがあるので，5点もらえて，そこでエピソード終了になります。

- **Dockerを利用している人は，`C`が出てこない人がいるかも知れません。その場合は西井に連絡してください。**

ここから下には，自由にセルを作成して，いろいろと動作確認して構いません。

- ここまでに作成したエージェントを使って，`MDPGridworld-v3`の攻略をしてみてください。

```python

```