<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

# リスト(list), p.141

## リスト/タプル/辞書/セット

| 名前       |    例                       |                                                      |
|-------------|----------------------------|------------------------------------------------------|
| **リスト(list)**   | ``L=["kiwi", "banana", "apple"]``  | **インデックス番号付き集合**               |
| タプル(tuple)| ``T=("kiwi", "banana", "apple")``      |インデックス番号付き集合(代入不可)       |
| 辞書(dict)  | ``D={"kiwi":3, "banana":2, "apple":3}`` | インデックス番号無し (key,value) 対応表 |
| セット(set)   | ``S={"kiwi", "banana", "apple"}``     | インデックス番号無し集合                |

## リストの定義, p.142

### 要素を列挙して定義

```python
L=[2,5,8,3]
L2=["cat","dog"]
```

リストの内容は`print()`で表示できる。

```python
print(L)
print(L2)
```

```python
L=[1,2,"cat",["banana","kiwi"]]
print(L)
```

### range()を使って定義

以前の課題(basic)に登場した関数`range()`の**返り値はリストではない**が，`range()`が生成する数列を関数`list()`でリスト化できる。

```python
L=list(range(5))
print(L)
```

関数`range()`の返り値がリストでないことは以下を実行すると確認できる。

```python
val_range=range(5)
print(val_range)
print(type(val_range))
```

Q. 関数`range()`と関数`list()`を使って，0以上20以下の**偶数**を要素とするリスト`L`をつくりなさい。for文等は使わず，1行で命令文を書くこと。

<!-- grade_id: cell-bbf22f37b1a397a5, points: 0, grade: false, solution: true -->
```python
# L= 右辺をrange()を使って作り，以下の解答欄に記入
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Lの値を確認しなさい。
print(L)
```

<!-- grade_id: cell-18e3cf3ea8a85310, points: 1, grade: true, solution: false -->
```python

```

Q. **リストの要素数は関数`len()`で判定**できる。前問のリスト`L`の要素数を`len()`を使って求め，変数`ans`に代入しなさい。
```

<!-- grade_id: cell-3c317238ab83dae2, points: 0, grade: false, solution: true -->
```python
# 解答方法) ans=  右辺にlen()を使ってリストLの要素数を調べる命令文を書く
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: cell-adb60c0361b980c7, points: 1, grade: true, solution: false -->
```python
#評価用セル
```

## リストの乗算

リストの乗算を使うと，簡単に要素数を増やすことができる。

```python
L=["sheep"]*5
print(L)
```

```python
# このセルで試してみよう
```

Q. 1が10個入っているリスト`x`を乗算記号`*`を使って作りなさい (`x=[1,1,1,1,1,1,1,1,1,1]`を生成)。

<!-- grade_id: cell-e0c7f2092e17d7f0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: cell-6ff221d923027344, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## リストの加算, p.159

Q. 下記の演算を実行すると，変数`v`や`fruit`の値はどうなるか予想して確認しなさい。

```python
x=[1,2,3,4]; y=[9,8,7,6]
v=x+y
fruit=[]
fruit+=["apple"]
fruit+=["kiwi"]
z=x+fruit
# print()を使って，各変数値を確認すること
```

Q. 以下の仕様を満たす関数`bind_repeat(a,b,n)`を作りなさい。

- 概要: 2つのリスト`a`, `b`を結合し，その要素を`n`回反復したリストを作る
- 引数
    - `a` : リスト
    - `b` : リスト
    - `n` : 反復する回数(デフォルト値1)
- 返り値
    - リスト`a`, `b`の要素が交互に`n`回反復して現れるリスト
- 注意
    - リストの加算と乗算を利用する。

**実行例**

- 命令文
```python
a=["0"]
b=["1","2"]
res=bind_repeat(a,b)
print(res)
res=bind_repeat(a,b,3)
print(res)
```
- 出力結果
```python
['0', '1', '2']
['0', '1', '2', '0', '1', '2', '0', '1', '2']
```

<!-- grade_id: cell-bc2c841844177171, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-1002cac6f929b4f1, points: 1, grade: true, solution: false -->
```python
# 評価用セル
eq_(bind_repeat(["0"],["1"]),['0', '1'])
eq_(bind_repeat(["0"],["1"],3),['0', '1', '0', '1', '0', '1'])
eq_(bind_repeat(["cat"],["dog"],3),['cat', 'dog', 'cat', 'dog', 'cat', 'dog'])
print("ok")
```

## リストの参照とスライス, p.146, p.160

- リストの一部分を取り出す操作のことを**スライス**とよぶ。
- 要素の**インデックス**を指定することで，スライスできる。
- **注意**: リストのインデックスは，C言語と同様に0から開始。

### スライス範囲の指定方法

リスト`L`に対するスライス方法

- `L[start]` : インデックスが`start`以上
- `L[start:end]` : インデックスが`start`以上`end`**未満**
- `L[start:end:step]` : インデックスが`start`以上`end`**未満**，ただし，切り出す要素のインデックスはstepずつ増加
- **範囲指定に負の値を用いて，リストの末尾からの指定も可能**

例)
```python
L=["apple", "banana", "strawberry", "orange", "grape"]
a=L[0]
b=L[-1]
c=L[-2]
d=L[0:3] # インデックスが0から3未満(以下ではない!)の部分の切り出し
e=L[:2]  # l[0:2]と同じ
f=L[-2:] # l[-2:] # インデックス-2から最後まで
print(a,b,c,d,e,f,sep="\n")
```
実行結果は以下のようになる。
```python
apple
grape
orange
['apple', 'banana', 'strawberry']
['apple', 'banana']
['orange', 'grape']
```

```python
# このセルで試してみよう
```

Q. 以下の仕様を満たす関数`head(L,n)`を定義しなさい。

- 概要: リストの先頭部分の要素のみ取り出す
- 引数
    - `L` : リスト
    - `n` : 表示する要素の個数(デフォルト値を3とする)
- 返り値
    - リスト`L`の先頭から`n`個の要素のリスト
- 注意
    - `for`文や`whild`文は使わないで，リストのスライスにより処理する

**実行例**

- 命令文
```python
L=["apple", "banana", "strawberry", "orange", "grape"]
res=head(L)
print(res)
```
- 出力結果
```python
["apple", "banana", "strawberry"]
```

<!-- grade_id: cell-ac7e9c82065761c2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-3a79eab96a961b44, points: 1, grade: true, solution: false -->
```python
# 評価用セル
test_list=list(range(10))
eq_(head(test_list),[0,1,2])
eq_(head(test_list,2),[0,1])
test_list2=["apple", "banana", "strawberry", "orange", "grape"]
eq_(head(L=test_list2,n=4),["apple", "banana", "strawberry","orange"])
print("ok")
```

Q. 以下の仕様を満たす関数`tail(L,n)`を定義しなさい。

- 概要: リストの最後の部分の要素のみ取り出す
- 引数:
    - `L` : リスト
    - `n` : 数値(デフォルト値を3とする)
- 返り値
    - リスト`L`の要素のうち，最後から`n`個分の要素のリスト
- 注意
    - `for`文や`whild`文は使わないで，リストのスライスにより処理する

**実行例**

- 命令文
```python
L=["apple", "banana", "strawberry", "orange", "grape"]
res=tail(L)
print(res)
```
- 出力結果
```python
["strawberry", "orange", "grape"]
```

<!-- grade_id: cell-1e0157132abd8094, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-dcf06769c0abd529, points: 1, grade: true, solution: false -->
```python
# 評価用セル
test_list=list(range(10))
eq_(tail(test_list),[7,8,9])
eq_(tail(test_list,2),[8,9])
test_list2=["apple", "banana", "strawberry", "orange", "grape"]
eq_(tail(L=test_list2,n=4),["banana", "strawberry", "orange", "grape"])
print("ok")
```

Q. 以下の仕様を満たす関数`from_to(L,a,b)`を定義しなさい。

- 概要: リストのうち，指定範囲の要素のみ取り出す
- 引数
    - `L` : リスト
    - `a` : 取り出す要素範囲の開始インデックス
    - `b` : 取り出す要素範囲の終了インデックス
- 返り値
    - リスト`L`の要素のうち，第`a`要素から，第`b`要素まで($a\le i\le b$を満たす全ての`L[i]`)のリスト
- 注意
    - リストの要素は第0要素から数える
    - `for`文や`whild`文は使わないで，リストのスライスにより処理する
    - `a`, `b`が0以上の値の場合に正常動作すれば良い

**実行例**

- 命令文
```python
L=["apple", "banana", "strawberry", "orange", "grape"]
from_to(L,2,3)
```
- 出力結果
```python
["strawberry", "orange"]
```

**実行例2**

- 命令文
```python
L=["apple", "banana", "strawberry", "orange", "grape"]
from_to(L,-3,-2)
```
- 出力結果
```python
["strawberry", "orange"]
```

<!-- grade_id: cell-32f1bacfe75bcd86, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
list(range(10))
```

<!-- grade_id: cell-8ee74ce8ccb4ce0c, points: 1, grade: true, solution: false -->
```python
# 評価用セル
test_list=list(range(10))
eq_(from_to(test_list,3,4),[3,4])
eq_(from_to(L=test_list,a=1,b=2),[1,2])

test_list2=["apple", "banana", "strawberry", "orange", "grape"]
eq_(from_to(test_list2,1,3),["banana", "strawberry", "orange"])
eq_(from_to(test_list2,0,0),["apple"])
eq_(from_to(L=test_list2,a=0,b=0),["apple"])
print("ok")
```

### スライス範囲の指定方法その2

リスト`L`の要素を間引いてスライスすることもできる。

`L[start:end:step]`
- インデックスが`start`以上`end`**未満**の要素をスライス。
- ただし，取り出す要素のインデックスは`step`ずつ増加

例)
```python
L=["apple", "banana", "strawberry", "orange", "grape"]
a=L[::2]   # L[0:len(L):2] と同じ意味
b=L[::-1]  # L[0:len(L):-1] と同じ意味
print(a,b,sep="\n")
```
実行結果
```python
['apple', 'strawberry', 'grape']
['grape', 'orange', 'strawberry', 'banana', 'apple']
```

```python
# このセルで試してみよう
```

Q. 以下の仕様を満たす関数`odd(L)`を定義しなさい。

- 概要: リストから，インデックスが1,3,5...の要素のみを取り出す。
- 引数
    - `L` : リスト
- 返り値
    - リスト`L`の要素のうち，奇数番目のみを要素とするリスト
- 注意
    - `for`文や`whild`文は使わないで，リストのスライスにより処理する

**実行例**

- 命令文
```python
L=["apple", "banana", "strawberry", "orange", "grape"]
res=odd(L)
print(res)
```
- 出力結果
```python
["banana", "orange"]
```

<!-- grade_id: cell-ae7c3db225080e76, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-823543d3c828d919, points: 1, grade: true, solution: false -->
```python
# 評価用セル
test_list=list(range(10))
eq_(odd(test_list),[1,3,5,7,9])
test_list2=["apple", "banana", "strawberry", "orange", "grape"]
eq_(odd(L=test_list2),["banana","orange"])
print("ok")
```

### リストの要素とインデックスの参照
リストの要素をインデックスとともに参照する方法の例
```python
L=[2,5,8,3]
for i in range(len(L)):
    print("L[%d]=%d" % (i,L[i]))
```
リスト内の各要素のインデックスと値の両方を返す関数`enumerate()`を用いて，以下のように書くこともできる
```python
for i,x in enumerate(L):     # 関数enumerate()で各要素のインデックスiと値L[i]に順にアクセス
    print("L[%d]=%d" % (i,x))
```

Q. for文と関数`enumerate()`を使って，`L=[1,"apple",[2,3],3.5]` の各要素とそのタイプ(型)を順に，以下の例のように表示する関数`show_list()`を作りなさい。

**実行例**

- 命令文
```python
L=[1,"apple",[2,3],3.5]
show_list(L)
```
- 出力結果
```python
L[0]=1, <class 'int'>
L[1]=apple, <class 'str'>
L[2]=[2, 3], <class 'list'>
L[3]=3.5, <class 'float'>
```

<!-- grade_id: cell-c1aa45ac7f0ec1b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-ec19f8c696b24755, points: 1, grade: true, solution: false -->
```python
# 評価用セル
test_list=[1,"apple",1.0]
with capture_stdout() as capture:
        show_list(test_list)
res=capture.result
assert res.replace(" ","").strip() =="L[0]=1,<class'int'>\nL[1]=apple,<class'str'>\nL[2]=1.0,<class'float'>"

test_list=[(1,3)]
with capture_stdout() as capture:
        show_list(L=test_list)
res=capture.result
assert res.replace(" ","").strip() =="L[0]=(1,3),<class'tuple'>"

# check if enumerate is used
from unittest.mock import patch
with patch('__main__.enumerate') as mock_enumerate:
    show_list(test_list)

mock_enumerate.assert_called_once_with(test_list)
print("ok")
```

## リストの編集メソッド, p.151
- `list.append(x)`    # 末尾に要素`x`を追加
- `list.extend(list)` # 末尾に`list`の要素を追加
- `list.insert(i,x)` # 位置`i`に`x`を挿入
- `list.remove(i)`   # 位置`i`の要素を削除
- `list.clear()`     # リスト内の要素をすべて削除

実行例
```python
fruit=[]

fruit.append("apple")
print(fruit)

fruit.append(["kiwi","orange"])
print(fruit)

fruit.extend(["peach","melon"])
print(fruit)

fruit.insert(2,"banana")
print(fruit)
```
出力結果
```python
['apple']
['apple', ['kiwi', 'orange']]
['apple', ['kiwi', 'orange'], 'peach', 'melon']
['apple', ['kiwi', 'orange'], 'banana', 'peach', 'melon']
```

Q. 要素を含まない空リスト`to_buy`を作りなさい。

<!-- grade_id: cell-0dcb926e1ad6f26a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: cell-f646425f3d71328f, points: 1, grade: true, solution: false -->
```python
# 評価用セル
eq_(to_buy,[])
```

Q. 前問の`to_buy`に要素"cake"を追加しなさい。

<!-- grade_id: cell-60b5eb959c495df1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: cell-d167ae107ae1ebb5, points: 1, grade: true, solution: false -->
```python
# 評価用セル
eq_(to_buy,['cake'])
```

Q. 前問のリスト`to_buy`に，リスト`vegetables=["tomato","cabbage"]`を結合し，`to_buy`の最後の2つの要素が"tomato","cabbage"になるようにしなさい。

<!-- grade_id: cell-c80d7c5a4e9be62f, points: 0, grade: false, solution: true -->
```python
vegetables=["tomato","cabbage"]
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: cell-fce8ac150d37f283, points: 1, grade: true, solution: false -->
```python
# 評価用セル
eq_(to_buy,['cake', 'tomato', 'cabbage'])
```

Q. 前問のリスト`to_buy`に，リスト`beverages=["beer","wine"]`を追加し，`to_buy`の最後の要素が`["beer","wine"]`になるようにしなさい。(最後の要素が`"wine"`ではなく，リスト`["beer","wine"]`になることに注意。)

<!-- grade_id: cell-4c0bbd998f972622, points: 0, grade: false, solution: true -->
```python
beverages=["beer","wine"]
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: cell-39785ed92d8fa0be, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 前問のリスト`to_buy`の`['beer', 'wine']`の前に要素`"carrot"`を追加しなさい。

<!-- grade_id: cell-33000d115931c1a6, points: 0, grade: false, solution: true -->
```python
# 評価用セル
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: cell-25aaaa4833099b5a, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## リストの代入, p.165

変数への，リストオブジェクトの代入は**参照渡し**である。
- pythonではオブジェクトの生成時に固有の`id`がふられる。
- リストオブジェクトの代入の際は，`id`のみがコピーされるが，実体はコピーされない。

例)
```python
dat1=[1,2,3]  # リスト[1,2,3]のidがdat1の参照先に
dat2=dat1     # data1の参照先idをdat2にコピー
dat2[-1]=5
print(dat1)
```
出力結果
```python
[1, 2, 5]   # dat2の要素を変更したら，dat1の要素も変わっている...
```
`dat1`と`dat2`の`id`を確認してみる。
```
print(id(dat1))
print(id(dat2))
```
出力結果(`id`は実行するたびに異なる。)
```python
140234586538440
140234586538440
```
同じ`id`になっていることがわかる。

```python
# このセルで試してみよう
```

## オブジェクトの比較
- `x is y`
    - `x` と `y` の`id`が同じ場合に`True`
- `x == y``
    - `x` と `y` の値が同じときに`True`
    - この値が`True`でも，`id`が同じとは限らない
    
**実行例**
```python
dat1=[1,2,3]
dat2=dat1
print(dat1 is dat2)
print(dat1 == dat2)
```
出力結果
```python
True
True
```

```python
# このセルで試してみよう
```

## リストの複製(値渡し), p.167

リストの(`id`の**参照渡しではなく，値渡し**による)コピーの方法。
- 方法1: メソッド`list.copy()`を使う
```python
dat2=dat.copy()
```
- 方法2: スライスを使う
```python
dat3=dat[:]   # dat1のはじめから終わりまで切り出す
```
- 方法3: リストを関数`list()`でリスト化する
```python
dat4=list(dat)
```

Q. 次の命令文の実行後の状態として，以下から正しいものを選びなさい。

```python
dat1=[1,2,3]
dat2=dat1
dat1=[5,6,7]
```
1. `dat1`と`dat2`の`id`は同じであり，`dat1`と`dat2`の各要素の値も同じである。
2. `dat1`と`dat2`の`id`は異なりが，`dat1`と`dat2`の各要素の値は同じである。
3. `dat1`と`dat2`の`id`は同じだが，`dat1`と`dat2`の各要素の値は異なる。
4. `dat1`と`dat2`の`id`は異なり，`dat1`と`dat2`の各要素の値も異なる。

解答欄には，以下の例のように正しい番号を変数`ans`の値として答えること。
```python
ans=5
```

```python
# このセルで試してみよう
```

<!-- grade_id: cell-afa8c5a830cd6267, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-711f7f2d066d1411, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 次の命令文の実行後の状態として，以下から正しいものを選びなさい。

```python
dat1=[1,2,3]
dat2=dat1.copy()
dat1=[5,6,7]
```
1. `dat1`と`dat2`の`id`は同じであり，`dat1`と`dat2`の各要素の値も同じである。
2. `dat1`と`dat2`の`id`は異なりが，`dat1`と`dat2`の各要素の値は同じである。
3. `dat1`と`dat2`の`id`は同じだが，`dat1`と`dat2`の各要素の値は異なる。
4. `dat1`と`dat2`の`id`は異なり，`dat1`と`dat2`の各要素の値も異なる。

解答欄には，以下の例のように正しい番号を変数`ans`の値として答えること。
```python
ans=5
```

```python
# このセルで試してみよう
```

<!-- grade_id: cell-abd656179c3b8eba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-164e3958dabf1035, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 次の命令文の実行後の状態として，以下から正しいものを選びなさい。

```python
dat1=[1,2,3]
dat2=dat1
dat1[1]=5
```
1. `dat1`と`dat2`の`id`は同じであり，`dat1`と`dat2`の各要素の値も同じである。
2. `dat1`と`dat2`の`id`は異なりが，`dat1`と`dat2`の各要素の値は同じである。
3. `dat1`と`dat2`の`id`は同じだが，`dat1`と`dat2`の各要素の値は異なる。
4. `dat1`と`dat2`の`id`は異なり，`dat1`と`dat2`の各要素の値も異なる。

解答欄には，以下の例のように正しい番号を変数`ans`の値として答えること。
```python
ans=5
```

```python
# このセルで試してみよう
```

<!-- grade_id: cell-695d5d0320b2d650, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a94fd1b8b01d1122, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## リストの並べ替え, p.168

### 関数`sorted()`
```python
sorted(list)
```
- 関数`sorted()`はオブジェクトは書き換えず，ソート後のリストを返す

### メソッド`list.sort()`

```python
list.sort(reverse=False) 
```
- 引数は省略可。降順にするときは `reverse=True`を指定
- `list.sort()`メソッドはオブジェクト`list`を書き換える

Q. `L=[1,6,3,7,2]`を，関数`sorted()`でソートし，その結果を変数`sorted_L`に代入しなさい。

<!-- grade_id: cell-c2565c06fc066d5e, points: 0, grade: false, solution: true -->
```python
L=[1,6,3,7,2]
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果の確認をしておきましょう
```

<!-- grade_id: cell-cf9aaa31071061bd, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. `L=[1,6,3,7,2]`をメソッド`list.sort()`でソートしなさい。

<!-- grade_id: cell-89b763a558605a86, points: 0, grade: false, solution: true -->
```python
L=[1,6,3,7,2]
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果の確認をしておきましょう
```

<!-- grade_id: cell-402e1b7f0c8f134c, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## リストの検索, p.177

### 検索メソッド
- `list.count(x)`       # 値がxである要素の数を取得
- `list.index(x)`       # 値がxである要素のインデックスのうち最小値を返す。

### 検索・計算関数
- `max(list)`
- `min(list)`
- `sum(list)`

Q. 数値リストに対して，最大値をとる要素のインデックスのうち最小のもの返す関数`argmax()`を，関数`max()`とメソッド`list.index()`を使って作りなさい。

引数
- `L`: リスト

返り値
- リスト`L`の要素のうち，値が最大になるもののインデックス。最大値をとる要素が複数あるときには，その中で最小のインデックスを返すものとする。

実行例
```python
numbers=[5,7,8,3,8,3]
index=argmax(L=numbers)
print(index)
```
出力結果
```
2
```

<!-- grade_id: cell-f15991962ed7fc2b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-c0cb3335d2c57f86, points: 1, grade: true, solution: false -->
```python
# 評価用セル
assert_equal(argmax([1,2,3,4]),3)
assert_equal(argmax(L=[3,4,4,4,4]),1)
print("ok")
```

### リスト要素のキーワード検索

**書式:** `<keyword> in <list>`
- リスト<list>の要素のうち，キーワードに一致するものがあるか否かを判定する
- 返り値: True or False
    
```python
to_buy=["apple", "banana", "strawberry", "orange", "grape"]

def should_i_buy(item):
    judge=item in to_buy
    print("Should I buy a", item, "?", judge)

should_i_buy("banana")
```

Q. 以下の仕様を満たす関数`check_action(action)`を作りなさい。

- 概要: 引数で与えられたオブジェクトが，大域変数(リスト)`action_list`に含まれているかどうかを確認する
- 引数名: `action`
- 返り値: 引数`action`が大域変数(リスト)`action_list`に含まれていたら，該当する要素のインデックス(`list.index()`を使う)を返し, 含まれていなければFalse

実行例
```python
action_list=["left","up","right","down"]
print(check_action("right"))
print(check_action("kiwi"))
```
出力結果
```python
2
False
```

<!-- grade_id: cell-677e9b6842ecdad4, points: 0, grade: false, solution: true -->
```python
action_list=["left","up","right","down"]
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-ea4c45a841440f01, points: 1, grade: true, solution: false -->
```python
# 評価用セル
action_list=["left","up","right","down"]
assert_equal(check_action("up"),1)
assert_equal(check_action(action="right"),2)
assert_equal(check_action("kiwi"),False)
action_list=["kiwi","banana","tomato"]
assert_equal(check_action("right"),False)
assert_equal(check_action("kiwi"),0)
print("ok")
```

## for文を用いたリスト生成

### リスト内包表記(list comprehension), p.173

目的: `for`文で生成される値をリスト化する方法

書式: **[ 式 for 変数 in イテラブル ]**
- リスト内包表記では，`for`の前に実行する式を書く
- **イテラブル**とは，1つずつ要素を取り出すことができるオブジェクトのこと。例えば，`list`, `tuple`, `dict`などや，`range()`などの関数。

例1) $1^2,2^2,3^2,\ldots,10^2$の値をリスト化する
```python
num10=[ i**2 for i in range(1,11) ]
```

例2) 1から6までの乱数を10回発生する。
```python
import random as ra
dice=[ ra.randint(1,6) for i in range(10) ]
```

Q. 以下の仕様を満たす関数`get_rand_action(n)`を作りなさい。

- 概要: 大域変数(リスト)`action_list`からランダムに要素を`n`個取り出す
- 引数名: `n`
- 返り値: 大域変数(リスト)`action_list`からランダムに要素を`n`個取り出してリスト化したもの
- 注意: ランダム選択には`random.choice()`を使うと便利(p.180)

実行例
```python
action_list=["left","up","right","down"]
print(get_rand_action(3))
```
出力結果(例)
```python
["up","down","right"]
```

<!-- grade_id: cell-366ad6ff1d38387d, points: 0, grade: false, solution: true -->
```python
import random as ra
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
action_list=["left","up","right","down"]
```

<!-- grade_id: cell-67ba45c1b50e71fe, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

### 条件付きリスト内包表記, p.175

目的: `for`文で生成される値のうち，**条件を満たす要素のみ**をリスト化する方法

書式: **[ 式 for 変数 in イテラブル if 条件式 ]**

例1) 30回乱数を生成する
```python
import random as ra
rnd=[ ra.random() for i in range(30) ]
```
例2) リスト`rnd`に含まれる0.1未満の数のみを抜き出す
```python
low01=[ i for i in rnd if i<0.1 ]
```
例3) 同上
```python
low01_2=[ x for i,x in enumerate(rnd) if x<0.1 ] # 同上
```

Q. 条件付きリスト内包表記を用いて，リスト`L`の要素のうち，指定した値`value`をとる要素のインデックスをリスト化して返す関数`index_list()`を定義しなさい。

引数
- `L`: リスト
- `value`: 値

返り値
- 引数のリスト`L`の要素のうち，値が`value`である要素のインデックスをリスト化して返す。

注意
- 処理には**関数`enumerate()`を利用すること**

**実行例1**)
```python
fruits = ["banana","kiwi","banana","apple"]
myans=index_list(fruits,"banana")
print(myans)
```
出力結果
```
[0, 2]
```

**実行例2**)
```python
Qlist = [2,4,1,4,3,2,4,1,4,3,4,1]
myans=index_list(L=Qlist,value=4)
print(myans)
```
出力結果
```python
[1, 3, 6, 8, 10]
```



<!-- grade_id: cell-3d22e74a12086122, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
Qlist=[2,4,1,4,3,2,4,1,4,3,4,1]
```

<!-- grade_id: cell-7f86569198c1c647, points: 2, grade: true, solution: false -->
```python
# 評価用セル
fruits = ["banana","kiwi","banana","apple"]
eq_(index_list(fruits,"banana"),[0, 2])
eq_(index_list(L=fruits,value="apple"),[3])
Qlist=[2,4,1,4,3,2,4,1,4,3,4,1]
eq_(index_list(Qlist,0),[])
eq_(index_list(Qlist,1),[2, 7, 11])
eq_(index_list(Qlist,4),[1,3,6,8,10])

from unittest.mock import patch
with patch('__main__.enumerate') as mock_enumerate:
    index_list(Qlist,4)
mock_enumerate.assert_called_once_with(Qlist)
print("ok")
```

## 2次元配列

2次元配列の定義の例
```python
L=[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]
```

Q. 以下を実行すると，どちらも各要素の値を0とする2次元リストができるが，どちらか一方，もしくはいずれの方法も使わないほうが良い。例えば，`L[2][1]=1`を実行することにより，不適切なものがいずれか，(もしくは両方か)を考えなさい。

方法1: 
```python
L= [[0]*3]*4    #リスト[0,0,0]を生成し，そのidを4回参照して2次元リストにする
```
方法2:
```python
L2= [[0 for i in range(3)] for j in range(4)]  # j が更新されるたびに，内側の[]内を実行
```

解答の方法: いずれか正しいと思うものを選んで，解答欄に書きなさい。
```python
ans=0 # いずれも不適切
ans=1 # 方法1のみが不適切
ans=2 # 方法2のみが不適切
```

<!-- grade_id: cell-f4d877e72904134e, points: 0, grade: false, solution: true -->
```python
# このセルで試してみよう
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d9bc8a8c1314242a, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```