<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">

# クラス(class), p.273

</div>

## 用語(クラス, インスタンス，オブジェクト)

### クラス
- pythonにおける型のようなもの。予め定義されているものもあるが，ユーザが定義して使うこともできる。
- 例) int, float, list, tuple, dict

### インスタンス
- クラス(型)から生成されるデータのこと。インスタンス・オブジェクトともいう
- intはクラス，1,2,..はインスタンス
- listはクラス, [1,3]はインスタンス

### オブジェクト(もの)
- 例) クラスもインスタンスもメソッドもオブジェクト。定義できるものはなんでもかんでもオブジェクト。。。

<div class="alert alert-success">

# クラスを作る, p.282

</div>

## キーワード
- **初期化メソッド**
    - 定義したクラスをインスタンス化するときに実行される関数
- **インスタンス変数**
    - クラス内の様々なメソッド(関数)で使うことができる変数。クラス内でのみ使える大域変数のようなもの。
    - 大抵，初期化メソッドで生成・参照。
    - クラス内では， `self.name`のように，必ず`self.`をつけて参照する(具体例は次のセルの問に)
    - インスタンス化したオブジェクトからは，`オブジェクト名.変数名` (`taro.name`など)で参照(具体例は次のセルの問に)
- **インスタンス・メソッド**
    - クラス定義中では第一引数に`self`を指定して定義する関数
    - クラス定義中では`self.関数名` で参照できる(引数からは`self`は省略)
    。インスタンス化したオブジェクトからは，`オブジェクト名.関数名` で参照(引数からは`self`は省略)

## 簡単な例

```python
class Name_list:
    def __init__(self, family="Momo", first="Taro"):  # 初期化メソッド(インスタンス生成時に実行)
        self.family = family  # インスタンス変数 family の生成
        self.first = first  # インスタンス変数には"self."をつける

    def print(self):  # インスタンス・メソッド(第一引数は必ずself)
        print(
            "My name is %s %s." % (self.first, self.family)
        )  # インスタンス変数には"self."をつけて参照
```

定義したクラスのインスタンス化や，インスタンス・メソッドの呼び出しの例。

```python
momo = Name_list()  # インスタンスを生成し，変数momoに代入
dai = Name_list(family="Yamaguchi", first="Daisuki")  # インスタンスを生成し，変数daiに代入

momo.print()  # インスタンスmomoの，インスタンス関数printを実行
dai.print()  # インスタンスdaiの，インスタンス関数printを実行
print(momo.family)  # インスタンスmomoの，インスタンス変数familyにアクセス
```

## 練習問題
**Q.** 以下のクラス宣言の間違いを修正し，クラス`Student`を完成しなさい。

```python
class Student:
    def __init__(self, name="someone"): # 初期化メソッド(インスタンス生成時に実行)
        self.name = name  # インスタンス変数 name の生成
        self.physics = 0  # インスタンス変数には"self."をつける
        self.math = 0     # インスタンス変数は，以下で定義するインスタンスメソッドのいずれでも使える

    def get_total(self):  # インスタンス・メソッド(第一引数は必ずself)
        total = self.physics+self.math  # 変数totalはget_score()のみで使う局所変数なので，インスタンス変数にはしない
        return(total)                 # (なんでもかんでもインスタンス変数にしてはいけない!)

    def print_total(): # インスタンス・メソッド。この行に間違いがある。
        print(name, self.get_total())  # この行にも間違いがある
```

<!-- grade_id: A-class-student, points: 0, grade: false, solution: true -->
```python
# 以下に修正後のクラスStudentを入力しなさい
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
taro = Student("Taro Yamada")  # class Studentのインスタンスを生成し，変数taroに代入
jiro = Student("Jiro Tanaka")  # class Studentのインスタンスを生成し，変数jiroに代入

taro.physics = 80  # インスタンス変数には，"オブジェクト名.インスタンス変数名"でアクセスできる(selfがオブジェクト名と置き換わる)
taro.math = 70

jiro.physics = 50
jiro.math = 90

taro.print_total()  # インスタンス・メソッドには，"オブジェクト名.メソッド名"でアクセスできる
jiro.print_total()

# 以上を実行したとき以下のように表示されたらOK
# Taro Yamada 150
# Jiro Tanaka 140

# Classの呼び出し方法を納得できるようにいろいろと試してみること。
```

<!-- grade_id: E-class-student, points: 2, grade: true, solution: false -->
```python
taro = Student("Taro Yamada")
jiro = Student("Jiro Tanaka")

taro.physics = 80
taro.math = 70
jiro.physics = 50
jiro.math = 90

with capture_stdout() as capture:
    taro.print_total()
    jiro.print_total()

res = capture.result
assert (
    res == "Taro Yamada 150\nJiro Tanaka 140\n"
)  # remove spaces, and spaces/tabs/\n at the end
```

<div class="alert alert-success">

# クラスの継承

</div>

```python
class Name_list:
    def __init__(self, family="Momo",first="Taro"):
        self.family = family
        self.first = first
    def print(self):
        print("My name is %s %s." % (self.first, self.family))
```
上記クラス`Name_List`(本シート初めの例と同じ)の改造版のクラス`Name_List_osk`を定義し，インスタンス・メソッド`print`の表示を日本語にしたかったとする。
このように，一部のインスタンス・メソッドの修正や追加のみをしたクラスを作成する場合，**クラスの継承**を用いると必要箇所のみの定義ですむ。

## スーパークラスの一部の関数を置き換える
以下の例では，初期化インスタンスは`class Name_list`と同じものを使うため再定義はせず，`print()`の再定義のみしている。

```python
class Name_list_osk(Name_list):  # class Name_Listを継承したクラス Name_list_jpを定義
    def print(self):
        print("うちの名前は%s%sやで。" % (self.family, self.first))
```

クラスAを継承してクラスBを定義した時，クラスAを**スーパークラス**，クラスBを**サブクラス**(もしくは派生クラス)と呼ぶ。上記の例では，`Name_list`がスーパークラス，`Name_list_osk`がサブクラスになる。親クラスと子クラスと呼ぶこともある。

以下は実行例。

```python
momo = Name_list()  # インスタンスを生成し，変数momoに代入
dai = Name_list_osk(family="山口", first="大鋤")  # インスタンスを生成し，変数daiに代入

dai.print()
```

## スーパークラスの関数を呼び出しつつ，さらに追加の処理も行う
以下は，クラス`Name_list`を継承し，学籍番号もインスタンス変数にもつクラス`NameID_list`を定義した例。
`super().__init__()`を呼び出すことで，スーパークラスの初期化メソッドを実行するとともに，追加の処理も行っている。

```python
class NameID_list(Name_list):
    def __init__(self, family="Momo", first="Taro", ID="000"):
        super().__init__(family, first)  # スーパークラス(Name_list)の初期化メソッドを呼び出す
        self.ID = ID

    def printID(self):
        print("ID:" + self.ID)
```

以下は実行例。スーパークラス(`Name_list`)の初期化メソッドも実行されているので，インスタンス変数`family`, `first`への値の代入もなされていることに注意。

```python
dai = NameID_list(family="Yamaguchi", first="Daisuki", ID="999")  # インスタンスを生成し，変数daiに代入
dai.printID()  # NameID_listで追加したインスタンス関数の呼び出し
dai.print()  # Name_listのインスタンス関数を呼び出せる
print(dai.family)  # スーパークラスの初期化メソッドにより，インスタンス変数`family`に値が代入されている。
```

## 練習問題
**Q.** 一つ前の練習問題のクラス`Student`を継承したサブクラス`Student3`を以下のように定義し，`english`の得点も追加した総合点を計算できるようにしなさい。

- 初期化インスタンスでは以下を実行
    1. `Student`クラスの初期化インスタンスの呼び出し
    2. インスタンス変数`english`の値を0に初期化
- クラス`Student`のインスタンス関数`get_total()`を定義し直して，インスタンス変数 `math`, `physics`, `english`の総得点を返り値にする。

<!-- grade_id: cell-7a0bd38c53187383, points: 0, grade: false, solution: true -->
```python
class Student3(Student):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-class-Student3, points: 2, grade: true, solution: false -->
```python
# 評価用セル
```