<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">

# お部屋の探検と強化学習

</div>

一列に並んだ`Nstate`個の部屋を強化学習エージェントが動き回り，総報酬を最大化する行動を学習するプログラムを作っいく。

環境:
- エージェントは一列に並んだ`Nstate`個の部屋を動き回ることができる。

状態: 
- エージェントは， 一列に並んだ`Nstate`個の部屋のどこ(`s=1,2...,Nstate`)にいるかを知ることができる

行動:
- 部屋の中にいるエージェントは行動として"left"か"right"を選ぶことができ，それぞれの行動により部屋番号が1つ小さい，もしくは1つ大きいほうの隣の部屋に移動する。
- 部屋1で"left"，もしくは部屋`Nstate`で"right"を選択すると，エピソードは終了となる。

報酬:
- 部屋`Nstate`で行動"right"を選択したときのみ報酬(`reward`)として10を得る
- それ以外の場合は報酬(reward)`-1`を得る


## 作成するプログラムの概要

以下は，これから作成していくクラスや関数と，その主な処理内容

`class Rooms`
- 環境を表現するクラス
- エージェントのとる行動を受け取り，その行動による状態遷移を決定する。また，遷移後の状態および獲得報酬を返したり，現在の状態を表示したりする。

`trial()`
- `Rooms`に対して，常に同じactionをとった場合の状態遷移をエピソード終了まで表示する
- `Rooms`の動作確認用のプログラム

`class Fool`
- エージェントのクラス
- 環境がどのような状態であっても，決まりきったactionや，ランダムactionのみを出力し，学習能力は持たない。
- 環境クラスのデバッグや，学習エージェクトとの性能比較に用いる。

`demo_env()`
- `Fool`と`Room`を接続し，行動選択と状態遷移のループを，指定したエピソード回数繰り返して実行する関数。

`class Qlearner`
- エージェントのクラス
- 環境の状態に応じた行動選択を，Q学習により行う

`play_env()`
- エージェント(`Fool`や`Qlearner`)と`Room`を接続し，行動選択と状態遷移のループを，指定したエピソード回数繰り返して実行する関数。
- `demo_env()`に，エージェントの学習機能を呼び出す命令を追加したもの

`plot_history()`
- `play_env()`を実行した結果出力される，報酬の変化の様子などをグラフ表示する関数

## 学習環境: class Rooms

**Q1.** 上述の説明にあったN部屋の環境を表現するクラス`Rooms`を定義しなさい。
    - 以下の仕様を満たしていれば，メソッドや変数の追加は自由にして良い
    - 注) 各メソッド内のみで使う変数はインスタンス変数にはしないこと!

**仕様**
- Class名: `Rooms`
- 概要
    - エージェントのとる行動を受け取り，その行動による状態遷移を決定する。また，遷移後の状態および獲得報酬を返したり，現在の状態を表示したりする。
- 初期化メソッド`__init__(self,Nstate=5)`
    - 引数は` Nstate` (部屋数)。そのデフォルト値は5
    - インスタンス変数の初期化
        - `state_list=[1,2,...,Nstate]`: 状態リスト(`range()`を使って定義)
        - `NState` : 部屋数。同名の引数の値を代入する(この変数は，他のインスタンスメソッドの実行に不要なら定義しなくても良い)
        - `r` :報酬値。初期値として-1を代入
- インスタンス・メソッド
    - `reset(self,s=None)`: エージェントの状態(どの部屋にいるか)をリセットする
        - 処理内容:
            - 引数`s`が`None`の場合，リスト `state_list` の要素からランダムに選択した値をインスタンス変数`s`(エージェントの状態)に代入(`random.choice()`を使う)
            - 引数`s`が`None`ではない場合，その引数値をインスタンス変数`s`(エージェントの状態)に代入
            - インスタンス変数`r`の値を-1に初期化
        - 返り値: インスタンス変数`s`
    - `step(self,a)`: エージェントの行動に応じた状態遷移をし，報酬決定やエピソード終了判定を行う
        - 処理内容:
            - 引数`a`(エージェントの行動)に応じて，インスタンス変数`s`(状態)を変更
            - 状態遷移による報酬値の決定と，エピソード終了判定
        - 返り値:
            - インスタンス変数`s`: 行動`a`による遷移後の新しい状態
                - エピソード終了時の状態`s`は，0 (状態1の左)もしくは`Nstate+1` (`Nstate`の右)とする
            - インスタンス変数`r`: 報酬値
            - 変数`terminal`: エピソード終了フラグ。エピソード終了時に`True`, それ以外には`False`
    - `render(self)`: エージェントがどこにいるかを，インスタンス変数`s`に応じて表示
        - 処理内容:
            - エージェントは@マークで表す。
            - 右外に，得られた報酬値を表示
            - 例) `s=2` の場合(`Nstate=5`の場合)
            ```
            |1|@|3|4|5| r=0
            ```
            - 例) `s=Nstate+1`の場合(`Nstate=5`の場合)
            ```
            |1|2|3|4|5|@ r=10
            ```
        - 返り値: なし

**その他**
`random`モジュールは以下のように宣言して使いなさい。
```python
import random as ra
```

<!-- grade_id: cell-7d4498210aa1d6ee, points: 0, grade: false, solution: true -->
```python
# プログラムは概ね以下のようになる
#
# class Rooms:
#    def __init__(....)
#        ...
#    def reset(...)
#        ...
#        return(self.s)
#
#    def step(self, a):
#        ...
#        return self.s, self.r, terminal
#
#    def render(self):
#        ...

import random as ra

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-rooms-Rooms-step, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
# test Rooms.reset()/step() for Nstate=5
rooms = Rooms(Nstate=5)

s = rooms.reset(s=5)
s, r, done = rooms.step("right")
eq_(s, 6)
eq_(rooms.s, 6)
eq_(r, 10)
eq_(rooms.r, 10)
eq_(done, True)
```

<!-- grade_id: E-rooms-Rooms-step2, points: 0.5, grade: true, solution: false -->
```python
# 評価用セルその2
# test Rooms.reset()/step() for Nstate=10
rooms = Rooms(Nstate=10)

s = rooms.reset(s=3)
eq_(s, 3)
s, r, done = rooms.step("right")
eq_(s, 4)
eq_(r, -1)
eq_(done, False)
s, r, done = rooms.step("left")
eq_(s, 3)
eq_(r, -1)
eq_(done, False)
```

<!-- grade_id: E-rooms-Rooms-step3, points: 1, grade: true, solution: false -->
```python
# 評価用セルその3
# test Rooms.reset() (random states)
result = [rooms.reset() for i in range(10)]
result2 = [rooms.reset() for i in range(10)]
assert_not_equal(result, result2)
```

<!-- grade_id: E-rooms-Rooms-render, points: 1, grade: true, solution: false -->
```python
# 評価用セルその4
# test Rooms.render()
with capture_stdout() as capture:
    rooms.reset(s=1)
    rooms.render()
    rooms.step("left")
    rooms.render()

res = capture.result
ans = " | @| 2| 3| 4| 5| 6| 7| 8| 9| 10|  r= -1\n\
 @| 1| 2| 3| 4| 5| 6| 7| 8| 9| 10|  r= -1"
assert res.replace(" ", "").strip() == ans.replace(" ", "").strip()
```

<!-- grade_id: E-rooms-Rooms-step10, points: 1, grade: true, solution: false -->
```python
# 評価用セルその5
```

## class Roomsのデバッグ用プログラム `trial()`
**Q2.** 前問で作ったクラス`Rooms`のデバッグのため，以下をみたす関数`trial()`を定義し，`Rooms`の各インスタンス・メソッドの動作が適切であることを確認しなさい。

- 引数
    - `a`: デフォルト値は"right"
    - `Nstate`: 状態数(部屋数)。デフォルト値は5
- 処理
    - 引数`a="right"`の場合は状態`s=1`から，`a="left"`の場合は状態`s=Nstate`から常に同じ行動`a`を用いて，エピソード終了まで状態遷移を繰り返す。状態の初期設定には，前問の`Rooms.reset()`を用いる。
    - 状態遷移には前問で作った`Rooms.step()`を用いる。
    - 状態遷移における各ステップの様子は前問で作った`Rooms.render()`を用いて表示する。`Rooms.step()`の返り値である報酬値の値も表示する。
    - 関数の終了は，エピソード終了時
- 注意
    - 以下の実行例と**全く同じ結果**が出力されたら合格! **必要に応じて前問の`Rooms`も修正**しなさい。
    

**実行例1**

- 命令文

```python
trial()
```

- 実行結果
```python
  | @| 2| 3| 4| 5|  r= -1
  | 1| @| 3| 4| 5|  r= -1
  | 1| 2| @| 4| 5|  r= -1
  | 1| 2| 3| @| 5|  r= -1
  | 1| 2| 3| 4| @|  r= -1
  | 1| 2| 3| 4| 5|@  r= 10
```

**実行例2**

- 命令文

```python
trial("left")
```

- 実行結果
```python
  | 1| 2| 3| 4| @|  r= -1
  | 1| 2| 3| @| 5|  r= -1
  | 1| 2| @| 4| 5|  r= -1
  | 1| @| 3| 4| 5|  r= -1
  | @| 2| 3| 4| 5|  r= -1
 @| 1| 2| 3| 4| 5|  r= -1
```

<!-- grade_id: cell-557a5e0c988b6706, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-rooms-trial-right, points: 1, grade: true, solution: false -->
```python
# 評価用セル
with capture_stdout() as capture:
    trial()

res = capture.result
ans = " |@| 2| 3| 4| 5| r=-1\n\
  | 1| @| 3| 4| 5|  r= -1\n\
  | 1| 2| @| 4| 5|  r= -1\n\
  | 1| 2| 3| @| 5|  r= -1\n\
  | 1| 2| 3| 4| @|  r= -1\n\
  | 1| 2| 3| 4| 5|@  r= 10"
assert res.replace(" ", "").strip() == ans.replace(" ", "").strip()
```

<!-- grade_id: E-roomts-trial-left, points: 1, grade: true, solution: false -->
```python
# 評価用セルその2
```

## 学習しないエージェント: `class Fool`
**Q3.** 以下のように，状態`s`がなんであろうとランダムに`action`の選択を行うか， 指定した`action`のみを選択する，学習しないエージェントのクラスを作りなさい。

- このクラスには，一見無駄な引数が存在するが，今後構築する環境クラスのテストや，強化学習エージェントとの性能比較に使うために必要なものである。

**仕様**
- クラス名: `Fool`
- 初期化メソッド
    - 引数とデフォルト値
        - `action_list`: とりうる行動の一覧のリスト
        - `action` : `Fool`エージェントがとる行動。デフォルト値は"random"(文字列)
    - 処理: 
        - 引数`action_list`の値を，インスタンス変数`action_list`に代入
        - インスタンス変数`a`の値を以下のように決定する
            - 引数`action`の値が引数(リスト)`action_list`の要素内にあるときは，`action`の値
            - 上記以外の場合の値は"random"
- インスタンス・メソッド
    - `get_action(self, s)`
        - 引数: `state`
        - 返り値: インタンス変数`a`の値によって以下のように決める
            - `a="random"`の場合はインスタンス変数`action_list`からランダム選択
            - 上記以外の場合は`a`の値
        - 注意: 返り値の決定に引数`s`は使わないが，今後作るエージェントのクラスとの互換性のために，あえて引数を設けている。
        
**実行例その1**
- 命令文
```python
action_list=["left","right"]
fool_rand=Fool(action_list)
print(fool_rand.get_action(s=1)) # sの値はなんでも良い
```
- 実行結果
```python
right # 実行ごとにランダムにかわる
```

**実行例その2**
- 命令文
```python
action_list=["left","right"]
fool_rand=Fool(action_list,"left")
print(fool_rand.get_action(s=1))
```
- 実行結果
```python
left # いつもこの値
```

<!-- grade_id: cell-4f90E-roomt-trial13618690a1cb6, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# プログラムは概ね以下のようになる
# class Fool:
#    def __init__(self,action_list,action="random"):
#        ...
#
#    def get_action(self,s):
#        ...
#        return(ret) # 選択したactionを返す
#
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-rooms-Fool, points: 1, grade: true, solution: false -->
```python
# 評価用セル
action_list = ["left", "right"]
fool = Fool(action_list)

result = [fool.get_action(s=1) for i in range(20)]
result2 = [fool.get_action(s=1) for i in range(20)]
assert_not_equal(result, result2)
```

<!-- grade_id: E-rooms-Fool2, points: 1, grade: true, solution: false -->
```python
# 評価用セルその2
fool = Fool(action_list, "left")
result = [fool.get_action(s=1) for i in range(20)]
assert "left" in result
assert "right" not in result

fool = Fool(action_list, "right")
result = [fool.get_action(s=1) for i in range(20)]
assert "right" in result
assert "left" not in result
```

## Foolの探検: `demo_env()`

**Q4.** 以下は，作成した環境`Rooms`を`Fool`エージェントが探検した結果や，その過程を表示する関数`demo_env()`の仕様である。

**引数とデフォルト値**
- `env`: 学習する環境
- `agent`: 環境を動く行動を決定するエージェント
- `s0` :初期状態(デフォルト値`None`)
- `max_episode`: 繰り返すエピソード数 (デフォルト値30)
- `max_step`: 1エピソードあたりの最大タイムステップ数 (デフォルト値30)
- `verbose`: 値が`True`のときに限り，以下の`render_step`で指定するエピソード毎に`Rooms.render()`を用いた状態遷移表示をする。(`render_step=10`ならば，第0,10,20,...エピソードの状態遷移を表示する)
- `render_step`: 何エピソードごとに状態遷移を表示するかを設定する変数 (デフォルト値20)
        
**返り値**
- `history_step`: 各エピソードにおける，エピソード終了までのステップ数のリスト
- `history_score`: 各エピソードにおいて獲得した総報酬のリスト

コメント文を参考にして，プログラムを完成しなさい。

```python
def demo_env(env, agent, s0=None, max_episode=30, max_step=30, verbose=False, render_step=10):
    # ここに，要素数がmax_episodeであるリストhistory_stepとhistory_scoreを定義する。
    # いずれも各要素の初期値は0とする

    # 各エピソードの総報酬は変数scoreに格納する。そのための命令文は以下のどこかに追加する
    for n_episode in range(max_episode):
        s = env.reset(s0)
        for t in range(max_step):
            # ここに，verbose==Trueの場合かつ，render_step回のエピソード実行毎に
            # env.render()を実行する命令文を書く
            action = agent.get_action(s)
            s,r,terminal = env.step(action)
            if terminal: break # <= if文を1行で書いてみた

        if verbose: 
            # ここに，render_step回のエピソード実行毎((n_episode%render_step==0 の時)に
            # env.render()を実行する命令文を書く
            print("Episode {} finished after {} timesteps with {} return".format(n_episode,t+1,score)) #時間tの表示は0からではなく，1からにする

        # history_step の第n_episode要素に，エピソード終了時のt+1を代入する命令文を追加
        # history_score の第n_episode要素に，エピソード終了時のscoreを代入する命令文を追加

    return(history_step,history_score)
```
実行例は後述

<!-- grade_id: cell-c0cc0b014f8694d6, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

以下の実行例を参考にデバッグしなさい。

**実行例**

- 命令文
```python
action_list = ["left","right"]
env = Rooms()
right_agent = Fool(action_list, action="right")
step,score = demo_env(env, right_agent, max_episode=3, verbose=True, render_step=1)
```
- 出力結果(実行するごとに結果は変わる)
```python
  | 1| 2| 3| @| 5|  r= -1
  | 1| 2| 3| 4| @|  r= -1
  | 1| 2| 3| 4| 5|@  r= 10
Episode 0 finished after 2 timesteps with 9 return
  | 1| @| 3| 4| 5|  r= -1
  | 1| 2| @| 4| 5|  r= -1
  | 1| 2| 3| @| 5|  r= -1
  | 1| 2| 3| 4| @|  r= -1
  | 1| 2| 3| 4| 5|@  r= 10
Episode 1 finished after 4 timesteps with 7 return
  | 1| 2| 3| @| 5|  r= -1
  | 1| 2| 3| 4| @|  r= -1
  | 1| 2| 3| 4| 5|@  r= 10
Episode 2 finished after 2 timesteps with 9 return
```

**注意**: 以下によく気をつけて動作確認をすること。

1. エージェント位置が部屋の外に出た時点まで表示されていること
2. エピソードごとに表示されるタイムステップ数や総報酬値が適切であること

```python
# このセルで動作確認をしよう
```

**Q4-2**
以下の各項目通りにプログラムが動作するか，次のセルで確認しなさい。

1. 上記のように`right_agent=Fool(action_list,action="right")`と定義した場合，スタート位置が`s=1`となったときの結果は， Q3で`trial("right")`を実行した場合と同じ結果になっていることを確認しなさい。
2. `left_agent=Fool(action_list,action="left")`を用いて`demo_env()`を実行しなさい。この場合，スタート位置が`s=5`になったときには，Q3で`trial("left")`を実行した場合と同様の結果になっていることを確認しなさい。
3. `fool_agent=Fool(action_list,action="random")`を用いた場合にも適切な行動遷移が得られることを確認しなさい。
4. `render_step=1`の場合は全てのエピソードについて， `render_step=2`の場合は， 第1, 3, 5, ...回目のエピソードの時に行動遷移が表示されることを確認しなさい。

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-rooms-demo_env-right, points: 1, grade: true, solution: false -->
```python
# 評価用セル
action_list = ["left", "right"]
env = Rooms()
right_agent = Fool(action_list, action="right")
left_agent = Fool(action_list, action="left")
random_agent = Fool(action_list)

# test 1 for Fool(right)
with capture_stdout() as capture:
    step, score = demo_env(
        env, right_agent, s0=5, max_episode=1, verbose=True, render_step=1
    )

res = capture.result
ans = "| 1| 2| 3| 4| @|  r= -1\n\
 | 1| 2| 3| 4| 5|@  r= 10\n\
Episode 0 finished after 1 timesteps with 10 return"
assert res.replace(" ", "").strip() == ans.replace(" ", "").strip()
```

<!-- grade_id: E-rooms-demo_env-left, points: 1, grade: true, solution: false -->
```python
# 評価用セル2
# test for Fool(left) and more
```

<!-- grade_id: E-rooms-demo_env-random, points: 1, grade: true, solution: false -->
```python
# 評価用セルその3
# test 3 for Fool(random)
```

## 学習エージェントクラス`Qlearner`の構築

**Q5.** クラス `Rooms`の環境に対して，報酬を得るための行動学習を行うエージェントのクラス`Qlearner`をつくりなさい。

- 以下の仕様を満たしていれば，メソッドや変数の追加は自由にして良い
- 注1) **各メソッド内のみで使う変数はインスタンス変数にはしないこと!**
- 注2) これまでの課題で作成した関数を使うと，かなりの部分が出来上がる
- 注3) 以前の課題(list課題)で作成した関数`index_list`を呼び出して使うと便利。


**仕様**
- クラス名: `Qlearner`
- 初期化メソッド
    - 引数とデフォルト値
        - `action_list` : 選択しうる行動のリスト
        - `epsilon=0.1` : random選択により行動決定をする確率(ε-greedy法の確率ε)
        - `alpha=0.2`   : learning rate
        - `gamma=0.9`: discount rate
        - `q0=0`          : Q値の初期値
    - インスタンス変数と初期化
        - `action_list` : 引数action_listの値で初期化
        - `epsilon` : 引数epsilonの値で初期化
        - `alpha`   : 引数alphaの値で初期化
        - `gamma`: 引数gammaの値で初期化
        - `q0`       : 引数q0の値で初期化
        - `Q={}`    : 空の辞書。キー`(s,a)`とQ値のペアを格納するために準備
- インスタンス・メソッド
    - `reset(self)`: 学習結果を初期化する
        - 実行内容: 辞書`Q`を空の辞書にする
    - `get_q(self, s, a)`
        - 引数: 
            - `s`: 状態
            - `a`: 行動
        - 辞書Qを参照し，`(s,a)`に対するQ値を返す
        - 辞書Qのキーに`(s,a)`がないときにはインスタンス変数`q0`の値を返す
    - `get_maxQ(self, s)`
        - 引数: 
            - `s`: 状態
        - 返り値: 状態`s`に対して，Q値が最大値をとる最適行動を選択した場合の以下の値
            - Q値
            - 最適行動の個数
            - 最適行動のインデックスのリスト。インデックスはリスト`action_list`中の要素順に対応
    - `get_action(self, s):`
        - 引数: 
            - `s`: 状態
        - 返り値: 状態`s`に基づいて以下のように決定した行動
            - 確率`epsilon`(インスタンス変数)でランダムな行動選択
            - Q値が最大となる行動が一意に決まるときは，その行動
            - Q値が最大となる行動が2つあるときは，ランダム選択
    - `learning(self, s1, a1, r, s2, terminal=False)`
        - 以下のQ学習則に基づいてQ値(インスタンス変数`Q[(s,a)]`)を更新する
        ```
        Q(s,a) <- Q(s,a) + alpha * (r + gamma*maxQ(s',a') - Q(s,a)) # 通常時
        Q(s,a) <- Q(s,a) + alpha * (r - Q(s,a)) # エピソード終了時
        ```
        - 引数:
            - `s1`: 学習則の`s`
            - `a1`: 学習則の`a`
            - `r`:  学習則の`r`
            - `s2`: 学習則の`s'`
            - `terminal=False`: エピソード終了状態かどうかを表す引数。
        - 学習則の計算について
            - 学習則の`alpha`, `gamma`はインスタンス変数を用いる
            - 学習則右辺の`Q(s,a)`には，インスタンス・メソッド`get_q()`を用いることができる
            - 学習則右辺の`max(Q(s'))`には，インスタンス・メソッド`get_maxQ()`の返り値を用いることができる
        - 返り値: なし

<!-- grade_id: A-rooms-indexlist, points: 0, grade: false, solution: true -->
```python
# class Qlearnerの実行に必要ならば，以前の課題(list課題)で作成したindex_list()をここに入力しなさい。
# 実行に不要なら，このセルには何も入力しなくて良い。
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: A-rooms-Qlearner, points: 0, grade: false, solution: true -->
```python
# このセルにclass Qlearnerの定義を入力しなさい
import random as ra

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-rooms-Qlearner-init, points: 1, grade: true, solution: false -->
```python
# 評価用セル
agent = Qlearner(action_list = ["up", "down"])
nt.assert_array_equal(agent.action_list, ["up", "down"])

# check init()
eq_(agent.epsilon, 0.1)
eq_(agent.alpha, 0.2)
eq_(agent.gamma, 0.9)
eq_(agent.q0, 0.0)
eq_(agent.Q, {})
```

<!-- grade_id: E-rooms-Qlearner-learn, points: 1, grade: true, solution: false -->
```python
# 評価用セル2
# check learning()
action_list = ["left", "right"]
agent = Qlearner(action_list)
agent.learning(s1=4, a1="right", r=-1, s2=5)
agent.learning(s1=5, a1="right", r=10, s2=6)
agent.learning(s1=3, a1="right", r=-1, s2=4)
agent.learning(s1=4, a1="right", r=-1, s2=5)
agent.learning(s1=5, a1="left", r=-1, s2=4)
agent.learning(s1=4, a1="right", r=-1, s2=5)
agent.learning(s1=5, a1="right", r=10, s2=6)
eq_(
    agent.Q,
    {
        (3, "right"): -0.2,
        (5, "left"): -0.2,
        (5, "right"): 3.6,
        (4, "right"): 0.16000000000000003,
    },
)
agent.learning(s1=4, a1="right", r=-1, s2=5)
agent.learning(s1=5, a1="right", r=10, s2=6)
```

<!-- grade_id: E-rooms-Qlearner-action, points: 1, grade: true, solution: false -->
```python
# 評価用セル3
# check get_action()
agent.epsilon = 0
actions = [agent.get_action(5) for i in range(10)]
assert "right" in actions
assert "left" not in actions
```

<!-- grade_id: E-rooms-Qlearner-reset, points: 1, grade: true, solution: false -->
```python
# 評価用セル4
# check reset()
agent.reset()
eq_(agent.Q, {})
```

## Qlearnerによる学習プログラム: `play_env()`

**Q6.** 環境`Rooms`において，上記で作成したクラス`Qlearner`による行動選択・学習を行う関数`play_env()`を作成する。既に作成した関数`demo_env()`を次のセルにコピー&ペーストし，さらに以下のように修正して，関数`play_env()`を作りなさい。また，これを用いてエピソードとともに学習が進むことを確認しなさい。

- **関数名**: `play_env()`に変更
- **引数**: `demo_env()`にあった引数の他に，**以下を追加**
    - `learning`: 
        - デフォルト値はTrue
        - 学習をするか否かのフラグ

- **動作**: 
    - 引数`learning==False`のときは，`demo_env()`と同じ。
    - 引数`learning==True`のときは，`demo_env()`の動作に加え，行動学習も行う
    

**動作確認**

```python
action_list = ["left","right"]
env = Rooms()
agent = Qlearner(action_list)
play_env(env,agent, learning=True, max_episode=30, verbose=True, render_step=10)
```
その他，引数の値を変えて，動作確認をする。エージェントのクラスにFoolを使うのも良いデバッグになる。

<!-- grade_id: A-rooms-play_env, points: 0, grade: false, solution: true -->
```python
def play_env(
    env,
    agent,
    s0=None,
    learning=True,
    max_episode=30,
    max_step=50,
    render_step=20,
    verbose=False,
):
    # 続きの処理を書く
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-rooms-play_env, points: 1, grade: true, solution: false -->
```python
# 評価用セル(Foolを用いた動作確認。demo_env()と同様の動作確認)

# Foolによるテスト(内容はdemo_envのテストと同じ)
action_list = ["left", "right"]
env = Rooms()
right_agent = Fool(action_list, action="right")
left_agent = Fool(action_list, action="left")
random_agent = Fool(action_list)

# test 1 for Fool(right)
with capture_stdout() as capture:
    step, score = play_env(
        env,
        right_agent,
        s0=5,
        max_episode=1,
        learning=False,
        verbose=True,
        render_step=1,
    )

res = capture.result
ans = "| 1| 2| 3| 4| @|  r= -1\n\
 | 1| 2| 3| 4| 5|@  r= 10\n\
Episode 0 finished after 1 timesteps with 10 return"
assert res.replace(" ", "").strip() == ans.replace(" ", "").strip()

# test 2 for Fool(left) and more
with capture_stdout() as capture:
    step, score = play_env(
        env,
        left_agent,
        s0=2,
        max_episode=3,
        learning=False,
        verbose=True,
        render_step=2,
    )

res = capture.result
ans = " | 1| @| 3| 4| 5|  r= -1\n\
  | @| 2| 3| 4| 5|  r= -1\n\
 @| 1| 2| 3| 4| 5|  r= -1\n\
Episode 0 finished after 2 timesteps with -2 return\n\
Episode 1 finished after 2 timesteps with -2 return\n\
  | 1| @| 3| 4| 5|  r= -1\n\
  | @| 2| 3| 4| 5|  r= -1\n\
 @| 1| 2| 3| 4| 5|  r= -1\n\
Episode 2 finished after 2 timesteps with -2 return"
assert res.replace(" ", "").strip() == ans.replace(" ", "").strip()
# test 3 for Fool(random)
```

<!-- grade_id: E-rooms-play_env-learning, points: 2, grade: true, solution: false -->
```python
# 評価用セルその2(agent.learning()の呼び出しと学習確認)
action_list = ["left", "right"]
env = Rooms()
agent = Qlearner(action_list)
agent.epsilon = 0
agent.Q = {(5, "right"): 4.0, (5, "left"): -2.0, (4, "right"): 2.0, (3, "right"): 0.2}
play_env(env, agent, s0=3, learning=True, max_episode=3, verbose=False, render_step=10)
eq_(
    agent.Q,
    {
        (5, "right"): 6.928,
        (4, "right"): 2.8544000000000005,
        (5, "left"): -2.0,
        (3, "right"): 0.5878400000000001,
    },
)
```

## パフォーマンスのグラフ表示: `plot_history()`

**Q7.** これまでに作成した`demo_env`や`play_env`の返り値である，各エピソード終了までのステップ数`step`と獲得総得点`score`の変化を**グラフ表示**するための命令文および関数`plot_history(step,score)`を作成しなさい。

- 以前の課題`matplotlib`のオブジェクト志向型インタフェースで作成する
- 2つのグラフは異なるグフフ上に描画すること
- 横軸はエピソード数
- 縦軸は，各エピソードの終了までのステップ数(`step`)および獲得総得点(`score`)
- `plot_history(step,score)`の返り値は，次セルに書いてあるように`ax1`,`ax2`とする。

`step`のグラフの注意
- y軸ラベル: "step"

`score`のグラフの注意
- x軸ラベル: "episode"
- y軸ラベル: "return"

共通の注意
- `step`や`score`に含まれる全エピソードをちょうど描画できるように描画範囲を設定
- x軸の刻み幅(tick)は，`step`に含まれるエピソード数を5で割った商に等しくする
- grid(格子)も表示する

`plot_history()`の呼び出し例
```python
action_list=["left","right"]
env=Rooms()
agent_fool=Fool(action_list)
step,score=play_env(env,agent_fool,learning=False) # Foolに学習メソッドの実装はしていないのでlearning=False
plot_history(step,score)
```
グラフの例
![](images/fool.png)

<!-- grade_id: cell-fe3bba4dd804928c, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt


def plot_history(step, score):
    fig = plt.figure(figsize=(7, 5))
    plt.rcParams["font.size"] = 16

    ax1 = fig.add_subplot(211)  # こちらにステップ数`step`を描画
    ax2 = fig.add_subplot(212)  # こちらに総報酬`score`を描画

    # YOUR CODE HERE
    raise NotImplementedError()

    return ax1, ax2
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-rooms-plot_history, points: 1, grade: true, solution: false -->
```python
# 評価用セル(ax1,ax2の共通設定確認)
```

<!-- grade_id: E-rooms-plot_history-label, points: 1, grade: true, solution: false -->
```python
# 評価用セルその2(ax1,ax2のラベル確認)
env = Rooms()
agent_fool = Fool(action_list=["left", "right"])
step, score = play_env(env, agent_fool, learning=False, max_episode=20)
ax1, ax2 = plot_history(step, score)
pc = PlotChecker(ax1)
pc.assert_ylabel_equal("step")

pc = PlotChecker(ax2)
pc.assert_xlabel_equal("episode")
pc.assert_ylabel_equal("return")
```

**Q7-2.**
Q学習エージェントによる学習パフォーマンスの変化も，前問と同様にグラフ表示しなさい。
学習エージェントの変数名は`agent_q`とすること。

グラフの例
![](images/q_agent.png)

<!-- grade_id: cell-2a3b3917e5190dca, points: 0, grade: false, solution: true -->
```python
# agent_q=... と，Qlearnerのインスタンスを生成し，play_env()による実行結果をplot_history()で表示する。
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-rooms-play_and_plot, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```