<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

##### このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py

<div class="alert alert-success">
    
# タプル(tuple)

</div>

| 名前       |    例                       |                                                      |
|-------------|----------------------------|------------------------------------------------------|
| リスト(list)   | ``L=["kiwi", "banana", "apple"]``  | インデックス番号付き集合               |
| **タプル(tuple)** | ``T=("kiwi", "banana", "apple")`` | **インデックス番号付き集合(代入不可)**    |
| 辞書(dict)  | ``D={"kiwi":3, "banana":2, "apple":3}`` | インデックス番号無し (key,value) 対応表 |
| セット(set)   | ``S={"kiwi", "banana", "apple"}``     | インデックス番号無し集合                |

- タプルの定義，要素へのアクセス，スライス等はリストとほぼ同様
- リストは要素を`[]`で囲むが，タプルは`()`で囲む
- リストと異なり，一度定義したら**要素の書き換えはできない**


## tupleを作る・使う, p.184-190

テキストを参考に以下を実行しなさい。

Q. "banana", "orange", "apple", "kiwi"を要素とするタプルfruitを定義しなさい。

<!-- grade_id: cell-f13fba9f69183491, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-tuple-fruit, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 前問で定義した`fruit`の最後の2つの要素をスライス(抜き出し)し，変数`to_buy`に代入しなさい。インデックスには負の整数を使うことでスライスすること。

<!-- grade_id: cell-89eba6f8eb151dd1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-tuple-slice, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 0から4までの整数を要素とするタプルを関数`range()`を使って定義し，変数`L`に代入しなさい

<!-- grade_id: cell-7f419a8ec32f2113, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-tuple-range, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 前問で定義した`L`の各要素を`a`,`b`,`c`,`d`,`e`の各変数に一括代入しなさい

<!-- grade_id: cell-67983b18d3d85e40, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-tuple-element, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
```

## タプル(tuple)への代入

一度定義したタプルの要素を後で書き換えることはできない。
```python
a=(1,2,3,5)
a[3]=4      # 要素を書き換えようとするとエラーが出る(試してみること)
```
ただし，以下のように変数に代入されているタプルを別のタプルに置き換えることはできる。
```python
a=(1,2,3,5)
a=(5,2,3)   # タプルの置き換え (タプル(5,2,3)のidがaの参照先に)
```
実際にを次セルで試してみなさい。

```python
# このセルで試してみよう
```

Q. 以下の命令文を実行したとき，エラーが出る行を全て答えなさい。ただし，各行頭の番号は行番号を表す。
```python
1:  a=[1,2,3,5]
2:  b=(9,8,7,5)
3:  a[3]=4
4:  b[3]=6
```
回答例)
```python
ans=[2]   # 2行目のみでエラーがでる場合
ans=[1,2]   # 1, 2行目でエラーがでる場合
```

<!-- grade_id: cell-de11f28de8559049, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-tuple-sustitution, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# 辞書(dict)

</div>

| 名前       |    例                       |                                                      |
|-------------|----------------------------|------------------------------------------------------|
| リスト(list)   | ``L=["kiwi", "banana", "apple"]``  | インデックス番号付き集合                 |
| タプル(tuple)| ``T=("kiwi", "banana", "apple")``      |インデックス番号付き集合(代入不可)       |
| **辞書(dict)**  | ``D={"kiwi":3, "banana":2, "apple":3}`` | **インデックス番号無し (key,value) 対応表** |
| セット(set)   | ``S={"kiwi", "banana", "apple"}``     | インデックス番号無し集合                |            |

- **書式:** `dict_name={key 1: value 1, key 2: value 2, ...}`
- 辞書(`dict`)には，検索のための`key`と，検索キーに紐付けられる`value`のペアを登録する。
- リストやタプルのように，各要素に対するインデックスは存在しない。

## 辞書(dict)を作る, p.214

以下は，`key`が文字列，`value`が数値の例

```python
reward={"cheese" : 100 , # リストや辞書等では，折返し記号"\"なしで改行可
        "enemy" : -10, 
        "friend" : -1}

print(reward["friend"]) # 辞書の参照は dict[key]
```

以下は，4-listにあったN部屋環境について，`key` を(状態(部屋番号), 行動)のタプル，`value`を，状態行動ペアに対する状態行動価値関数の値になっている例

```python
Qvalue={(1,"right") : 0.5, 
        (1,"left") : 1.5, 
        (2,"right") : 0, 
        (2, "left") : -1} 

print(Qvalue[(2,"right")])
```

Q. 以下の処理を次のセルに入力しなさい。
1. `key`を"alpha", "beta", "gamma", "delta", それぞれに対応する`value`を"a","b","r","d"とする辞書`abc`を定義しなさい。
2. 関数`len()`を使って辞書`abc`の要素数を確認し，変数`len_abc`に代入しなさい。
3. 辞書`abc`から`key` "beta"に対応する`value`を取得し，変数`value_beta`に代入しなさい。

<!-- grade_id: cell-6bb7508a653cc31b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: E-dict-basic, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## 辞書の更新・追加, p.218

辞書の`value`を変更したり，新しい`{key,value}`ペアを登録する方法

### dict[key]を用いる

書式: `dict[key]=new_value`
- `key`が辞書`dict`に存在しない場合は，`{key: value}`ペアを登録
- `key`が辞書`dict`に存在するは場合は，`key`に対応する`value`を`new value`に変更

```python
price={"tomato" : 200,
        "banana": 100, 
        "apple" : 150}

price["banana"]=120   # 辞書priceのkey "banana"に対するvalueを書き換える
print(price["banana"])

price["kiwi"]=200   # 辞書priceに無いkeyに対しては，{key,value}が登録される
print(price)
```

### dict.setdefault()を用いる

メソッド `dict.setdefault(key,value)`
- `key`が辞書`dict`に存在しない場合は，`{key: value}`ペアを登録
- `key`が辞書`dict`に存在するは場合は，何もしない
- 返り値は，`key`に対する`value`(上記処理後の値)

```python
print(price)
price.setdefault("banana",150)   # 辞書priceにkey "banana"がなければ登録する
print(price["banana"]) # すでに "banana" は登録済みだったので，valueの変更もされない
```

Q. 以下を実行した後の，`k`, `a`, `price["apple"]`の値を答えなさい。

```python
price={"tomato" : 200,
        "banana" : 180, 
        "apple" : 150}
k=price.setdefault("kiwi",250)
a=price.setdefault("apple",300)
```

解答欄には以下の例のように入力しなさい。
```python
val_k=500    # kの値
val_a=1000   # aの値
val_apple=30 # price["apple"]の値
```

<!-- grade_id: cell-df9f8502956bb5fe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-dict-setdefault, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## 辞書の参照, p.224

- 方法1: `dict[key]`
    - 直接辞書要素にアクセスする 
    - 存在しないキーに対してアクセスするとエラーになる
- 方法2: `dict.get(key)`
    - 存在しないキーに対してアクセスしてもエラーにはならず，`None`が返り値になる
- 方法3: `dict.get(key,[default])`
    - メソッド`dict.get()`は，**存在しないkeyにアクセスした場合に返す値(default)を指定**できる。
    - 注)`[]`で囲んだ引数は，省略可能なことを示す。

Q. 引数で与えられた辞書`D`に対して，引数で指定したリスト`key_list`の各要素(`key`)に対する`value`を表示する関数`check_key(D,key_list)`を作りなさい。

注意
- 辞書に存在しない`key`に対する`value`は，`dic.get()`の引数設定によって`999`と表示するようにしなさい。
- 表示結果は、下記の実行例のように`key`と`value`のペアを表示するようにすること。(`key`と`value`の間の区切り文字は":")

**実行例**

- 命令文
```python
price={"tomato" : 200 , "banana" : 100, "apple" : 150}
keys=["apple", "kiwi", "tomato"]
check_key(price,key)
```
- 出力結果
```python
apple:150
kiwi:999
tomato:200
```

<!-- grade_id: cell-f928f7728d467940, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-dict-checkkey, points: 1, grade: true, solution: false -->
```python
to_buy={"tomato" : 200 , "banana" : 100, "apple" : 150}
keys=["apple", "kiwi", "tomato"]

with capture_stdout() as capture:
    check_key(D=to_buy,key_list=keys)   # ここに，監視する関数を列挙する。
res=capture.result
assert res.replace(" ","").strip() =="apple:150\nkiwi:999\ntomato:200"

with capture_stdout() as capture:
    check_key({"potato" :300 , "pumpkin" : 200, "eggplant" : 100},["apple","eggplant"])   # ここに，監視する関数を列挙する。
res=capture.result
assert res.replace(" ","").strip() =="apple:999\neggplant:100"
print("Ok!")
```

## 辞書の値一覧, p.225

- `dict.keys()`   # 辞書に含まれるkey一覧を返す
- `dict.values()` # 辞書に含まれるvalue一覧を返す

Q. 以下の辞書`Q`のすべての`key`および`value`の一覧をそれぞれ取得し，それぞれ変数`key_list`, `value_list`に代入しなさい。

```python
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (2,"right") : 0, 
    (2, "left") : -1} 
```

<!-- grade_id: cell-4c1088cd7f1b0b5f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: E-dict-Q, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 引数で与えられた辞書`D`に対して，全ての`value`の総和を返り値とする関数`total_price(D)`を作りなさい。

**実行例**

- 命令文
```python
price={"tomato" : 200 , "banana" : 100, "apple" : 150}
sum=total_price(D)
print(sum)
```
- 出力結果
```python
450
```
ヒント) リスト要素の数値の合計は関数`sum()`で取得できる。

<!-- grade_id: cell-93aa327269732118, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: E-dict-total_price, points: 1, grade: true, solution: false -->
```python
# 評価用セル
fruit={"tomato" : 200 , "banana" : 100, "apple" : 150}
cost=total_price(D=fruit)
eq_(cost,450)
price={"potato" : 100 , "carrot" : 100, "egg" : 300}
cost=total_price(price)
eq_(cost,500)
print("ok")
```

## 関数とMutableな引数

関数の引数(仮引数)の値を関数内で変更した場合，
- 引数が**Immutable(変更不可)**な型ならば呼び出した側の引数(実引数)は変化しない(値渡し)
    - immutableな型の例: `int, float, str, tuple,...`
- **Mutable(変更可能)**ならば，実引数も変化する(参照渡し)
    - mutableな型の例: `list, dict, set,...`

Q. 下記のプログラムの実行後，`state`の値は変化しているだろうか。

```python
def new_state(s,a):
    if not a in action_list:
        print("invalid action")
        return(False)
    elif a=="right":
        s+=1
    else:
        s-=1
    return(s)

action_list=["right", "left"]
state=3
new_state(state,"right")
```

解答欄には以下のように記入しなさい。
```python
ans=False  #`state`の値が変化しない場合
ans=True   #`state`の値が変化する場合
```

<!-- grade_id: cell-e266d7697d5a3b22, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-dict-new_state, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 下記のように関数`append_action()`を作って，リスト`action_list`を更新しようとしてみた。
関数`append_action()`を呼び出した後，リスト`action_list` の要素は変化しているだろうか。

```python
def append_action(list,a):
        list.append(a)

action_list=["right", "left"]
append_action(action_list,"up")
```

```python
# このセルで試してみよう
```

解答欄には以下のように記入しなさい。

```python
ans=False  #`action_list`の要素が変化しない場合
ans=True   #`action_list`の要素が変化する場合
```

<!-- grade_id: cell-8ffe915e52638611, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-dict-append_action, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# 練習問題(Q学習のプログラムを作るための準備)

</div>

以下のように，辞書(dict) `Q` とリスト`action_list`がある。

```python
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (2,"right") : -1.0, 
    (3,"right") : 1.0, }
action_list=["right","left"]
```

`Q` は状態行動ペア`(s, a)`に対する状態行動価値関数の値(Q値)の一覧，`action_list`はとりうる行動のリストである。

Q1. 以下の仕様を満たす関数`get_maxQ(s,q0)`を作りなさい。
- **概要**
    - 引数で与えた状態`s`に対して, `Q(s,a)`の最大値 $\max_aQ(s,a)$ および， `Q(s,a)`の最大値を与える行動`a`, すなわち$\arg\max_aQ(s,a)$のリストを返す。
- **外部変数**(関数外に用意されている変数)
    - `Q`: 状態行動価値関数の値を格納する辞書(上記説明文参照)
    - `action_list`: エージェントが取りうる`action`のリスト(上記説明文参照)
- **引数**
    - s : 状態
    - q0 : 辞書`Q`に未定義なkey`(s,a)`に対するQ値。**デフォルト値は`0.0`**
- **返り値**: 引数の状態`s`に対して，Q値が最大値をとる最適行動を選択した場合の以下の2つの値
    - Q値
    - Q値を最大にする行動がリスト`action_list`のいずれの要素であるかを示す，そのインデックスのリスト
- **その他**
    - 最適行動を求める際，`Q`にない`(s,a)`ペアに対するQ値は`q0`とする。(以下の`Q`の例では`(2,"left")`, `(3,"left")`が未定義。`list.get()`の引数で処理する)
    - 以前の課題で作成してた関数`index_list()`も利用しなさい。
- **手順のヒント**
    1. 指定されたsに対し，`action_list`内の各行動とのペアに対するQ値の一覧リスト`qlist`を作る
    2. Q値の最大値`maxq`を求める
    3. `qlist`の要素のうち，値が`maxq`であるもののインデックスをリスト化する。(以前作成した関数`index_list()`を用いると簡単)
    
**実行例**

命令文
```python
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (2,"right") : -1.0, 
    (3,"right") : 1.0, }
action_list=["right","left"]

print(get_maxQ(1)) #状態がs=1の場合, Q値が最大
print(get_maxQ(2))
print(get_maxQ(s=3,q0=1.0))
```

出力結果
```python
(1.5, [1])
(0.0, [1])
(1.0, [0, 1])
```

<!-- grade_id: cell-c5638b73be930b32, points: 0, grade: false, solution: true -->
```python
# このセルに，以前の課題で作成した関数index_list()を記入しなさい。
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b2717978cdb01bc1, points: 0, grade: false, solution: true -->
```python
# このセルに関数get_maxQ()の定義を入力する
# ヒント) 概ね以下のようなプログラム構造になる
#
# def get_maxQ(s,q0=0.0):
#     qlist=[Q.get((s, a),q0) for a in ... 
#            
#     ...
#            
#    return maxq, argmaxq
#

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (2,"right") : -1.0, 
    (3,"right") : 1.0, 
  }

action_list=["right","left"]
print(get_maxQ(1))
# 他にもいろいろ試しなさい。
```

<!-- grade_id: E-dict-get_maxQ1, points: 1, grade: true, solution: false -->
```python
# 評価用セル
# 1st test
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (2,"right") : -1.0, 
    (3,"right") : 1.0, 
  }
action_list=["right","left"]

eq_(get_maxQ(1),(1.5, [1]))
eq_(get_maxQ(2),(0.0, [1]))
eq_(get_maxQ(s=3,q0=1.0),(1.0, [0,1]))
print("ok")
```

<!-- grade_id: E-dict-get_maxQ2, points: 1, grade: true, solution: false -->
```python
# 評価用セルその2
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (1,"up") : 1.0,
    (1,"down") : 1.0,
    (2,"right") : 1.0, 
    (2, "left") : 1.0,
    (2, "up") :   0.0,
  }
action_list=["right","left","up", "down"]

eq_(get_maxQ(1),(1.5, [1]))
eq_(get_maxQ(s=2),(1.0, [0, 1]))
```

Q2.以下の仕様を満たす関数`get_action(s,epsilon,q0)`を書きなさい。また，**動作確認**をしなさい。

- **概要**
    - 状態`s`における行動選択をε-greedy法によって行う。
- **外部変数**(関数外に用意されている変数)
    - `Q`: 状態行動価値関数の値を格納する辞書(前問参照)
    - `action_list`: エージェントが取りうる`action`のリスト(前問参照)
- **引数**
    - `s`: 状態
    - `epsilon` : ランダム選択の確率。**デフォルト値は0.1**
    - `q0` : 未定義なQ値に対する値。**デフォルト値は0.0**
- **返り値**: 状態`s`と， 辞書`Q`に登録されている要素に基づいて以下のように決定した行動
    - 確率`epsilon`で，とりうる全ての行動(リスト`action_list`の要素) の中からランダムに行動選択  (`random.choice()`を使う)
    - 上記以外の場合は**Q値(`value`)が最大となる行動**。
        - 辞書`Q`に定義されていない状態行動ペア`(s,a)`に対するQ値は`q0`として，最適行動を決定する。
        - **最大値をとる行動が複数あるときには， どれか一つをランダムに選択**する
- **その他**
    - 前問でつくった`get_maxQ()`を呼び出して使って良い
- **手順のヒント**
    1. `get_maxQ()`を使って，Q値が最大となるactionのインデックのリスト`argmaxq`を取得
    2. `argmaxq`からランダム選択したインデックスに対応する`action`が返り値。Q値が最大値となるactionが1個の場合もそうでない場合も，この方法で最適行動を選択できる(わからなければよく考えましょう)

**実行例**

命令文
```python
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (2,"right") : -1.0, 
    (3,"right") : 1.0, 
  }
action_list=["right","left"]

print(get_action(1,epsilon=0))
print(get_action(2,epsilon=0))
print(get_action(3,epsilon=0,q0=1.0))
print(get_action(1,epsilon=1.0))
```
実行結果
```python
left
left
right # この値はleftになることもある(確率1/2)
left  # この値はleftになることもある(確率1/2)
```

<!-- grade_id: cell-ad5b339903f2f96b, points: 0, grade: false, solution: true -->
```python
# ヒント) 概ね以下のようなプログラム構造になる
#
#def get_action(s, epsilon=0.1,q0=0.0):
#    if ra.random() < epsilon:
#        ret=               # 右辺を書く
#    else:
#        maxq, argmaxq=...   # このあとのプログラムを考える
#
#        ret=...
#       
#    return(ret)

import random as ra

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (2,"right") : -1.0, 
    (3,"right") : 1.0, 
  }
action_list=["right","left"]

print(get_action(1,epsilon=0))
# 他にもいろいろ試しなさい(状態や行動の種類が多い場合など)
```

<!-- grade_id: E-dict-get_action1, points: 1, grade: true, solution: false -->
```python
# 評価用セル
Q={(1,"right") : 0.5, 
    (1,"left") : 1.5, 
    (2,"right") : -1.0, 
    (3,"right") : 1.0, 
  }
action_list=["right","left"]
eq_(get_action(1,epsilon=0),"left")
eq_(get_action(2,epsilon=0),"left")
eq_(get_action(3,epsilon=0,q0=0),"right")

result=[get_action(3,epsilon=0,q0=1.0) for i in range(20)]
result2=[get_action(s=3,epsilon=0,q0=1.0) for i in range(20)]
assert_not_equal(result,result2)

result=[get_action(1,epsilon=1.0) for i in range(20)]
result2=[get_action(1,epsilon=1.0) for i in range(20)]
assert_not_equal(result,result2)
print("ok")
```

<!-- grade_id: E-dict-get_action2, points: 2, grade: true, solution: false -->
```python
# 評価用セルその2
```