<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">
   
# matplotlib, p.318

</div>

Pythonでグラフを描く方法はいろいろあるが，ここではmatplotlibを使った方法を紹介。
matplotlibには2種類の使い方がある。一つは**Pyplotインタフェース**，もうひとつは**オブジェクト指向インタフェース**を使う方法。単純なグラフを書くときには前者が便利だが，サブグラフを用意したり，細かい指定をするときには後者が便利。一応両方紹介。

まずはmatplotlibモジュールの読み込み。

```python
# 1行めは，matplotlibをjupyter notebookで使うための設定
%matplotlib inline
import matplotlib.pyplot as plt  # matplotlib.pyplotモジュールの読み込み
import math as ma
```

<div class="alert alert-success">

# Pyplot インタフェース
  
</div>  

sin関数とcos関数のグラフを一つのグラフに描いてみる。

```python
# 表示データ生成
step = list(range(200))  # データ生成用リスト(200点のデータを使ってグラフを生成する)
time = [0.1 * s for s in step]  # t軸データ
sint = [ma.sin(t) for t in time]  # y軸データ(リスト)その1(sin関数)
cost = [ma.cos(t) for t in time]  # y軸データ(リスト)その2(cos関数)

# グラフ描画
plt.figure(figsize=(8, 4))  # グラフサイズの指定(横:縦=8:4)
plt.plot(time, sint, label="sin t")  # x,y軸データから，グラフオブジェクトplt生成(labelは省略可)
plt.plot(time, cost, label="cos t")  # pltにcos関数のグラフも追加

# ラベル等の追加
plt.xlabel("time [s]")  # pltにx軸ラベル追加
plt.ylabel("y")  # pltにy軸ラベル追加
plt.legend(loc="upper right")  # pltに凡例(plot()の引数に指定したlabel)を追加

# plt.savefig('figure.png')    # グラフを保存したい場合。拡張子に応じてpng,jpg,pdf等の出力可能
plt.show()  # グラフオブジェクトpltの描画
```

## NumPyモジュール

数学関数の描画には，ベクトル・行列演算用の**NumPyモジュール**を使うと便利。NumPyはリストと似ているが数値データの処理に特化したメソッドがいろいろ用意されている。

上記プログラム例の，`sint`, `cost`のデータ生成部分は以下のように書き直せる。

```python
import numpy as np            # numpyモジュール読み込み

# 表示データ生成
step=0.1
time=np.arange(0,20+step,step)  # 0以上20+step未満の数値をstep刻みで生成
sint=np.sin(time)             # timeの各要素に対するsin関数の値をリスト化
cost=np.cos(time)             # timeの各要素に対するsin関数の値をリスト化

# グラフ描画 (ここから下は上のプログラムと同じ)
plt.figure(figsize=(8, 4))
# 以下省略...
```
上記プログラムを完成させて，次のセルで試してみなさい。

```python
# このセルで試してみよう
```

Q.`x=np.arange(0,1,0.1)`は，0以上1**未満**の数を0.1間隔で生成する。実際に実行し，どのようなオプジェクトが生成されたか`print()`文で確認しなさい。

<!-- grade_id: cell-434b88777753fba1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-mpl-nparrange, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 上記`x`に対し`y=x*x`を実行すると，どのような要素をもつオブジェクト`y`が生成されるか確認しなさい。

<!-- grade_id: cell-2d08f6e9bb344758, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-mpl-np-squared, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 上記`x`に対し`z=1-x`を実行すると，どのような要素をもつオブジェクト`z`が生成されるか確認しなさい。

<!-- grade_id: cell-99b91af0ad019a28, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-mpl-np-diff, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# オブジェクト指向インタフェース

</div>

## matplotlibの基本概念

- Primitives (基本パーツ)
    - 線や図形など，描画するオブジェクト。
    - 直線オブジェクト(Line2D), テキストオブジェクト(Text)などがある。
- [Axis](https://matplotlib.org/api/axis_api.html#matplotlib.axis.Axis) (座標軸オブジェクト)
    - ラベルや目盛りなどを要素とするの，座標軸を表現するクラス
    - x座標, y座標はそれぞれが別のAxisクラス
- [Axes](https://matplotlib.org/api/axes_api.html#matplotlib.axes.Axes) (座標コンテナ)
    - 座標軸(Axis)や描画グラフ(Primitives)を要素とするクラス
- [Figure](https://matplotlib.org/api/_as_gen/matplotlib.figure.Figure.html#matplotlib.figure.Figure) (図)
    - 全ての描画オブジェクトに対して，一番上の階層にあるクラス
    - Axesを要素とする。複数のAxesを登録にすれば，複数のグラフをきれいに配置して描画できる。


## 例

オブジェクト指向型のmatplotlibの使い方で，2つのグラフにsin関数とcos関数を描いてみる。

```python
# time, sint, costの部分は前述のNumpyを使った方法と同じなので省略(すでに実行済みの場合，再度定義の必要なし)

step = 0.1
time = np.arange(0, 20 + step, step)  # 0以上20+step未満の数値をstep刻みで生成
sint = np.sin(time)  # timeの各要素に対するsin関数の値をリスト化
cost = np.cos(time)  # timeの各要素に対するsin関数の値をリスト化

fig = plt.figure(figsize=(6, 5))  # Figure(図)を生成(横:縦=6:5)
plt.rcParams["font.size"] = 16  # グラフのタイトルなどのフォントの大きさをまとめて16pointに

ax1 = fig.add_subplot(211)  # Figureの中に座標(Axes)生成。2行1列に配置した1番目とする
ax2 = fig.add_subplot(212)  # 2行1列に配置した2番目のAxes生成

ax1.plot(time, sint, label="sin t")  # ax1にsin関数の折れ線グラフを登録(labelは省略可)
ax2.plot(time, cost, label="cos t")  # ax2にはcos関数

# ax1にラベル等の追加
ax1.legend(loc="upper right")  # plot()で指定したlabelを右上に配置
ax1.set_title("graph example")  # タイトル
ax1.set_xlabel("t [s]")  # x軸ラベル
ax1.set_ylabel("sin t")  # y軸ラベル
ax1.grid(True)  # 格子を描く
ax1.set_xlim(0, 20)  # 描画範囲は[0,5*pi]
ax1.set_ylim(-1, 1)  # 描画範囲は[-1,1]


# ax2にラベル等の追加
ax2.legend(loc="upper right")  # plot()で指定したlabelを右上に配置
ax2.set_xlabel("t [s]")  # x軸ラベル
ax2.set_ylabel("cos t")  # y軸ラベル
ax2.grid(True)  # 格子を描く
ax2.set_xlim(0, 20)  # 描画範囲は[0,5*pi]
ax2.set_ylim(-1, 1)  # 描画範囲は[-1,1]

plt.tight_layout()  # タイトルやラベルが重ならないように配置調整し，生成されているFigureを描画
```

一つのグラフにsinとcosの両方のグラフを描きたい場合は，上記の代わりに以下のようにすれば良い。
```python
ax=fig.add_subplot(111) # 1行1列に配置したグラフの1番目にAxes(座標)生成

ax.plot(time, sint, label="sin t") # axにsin関数の折れ線グラフを登録
ax.plot(time, cost, label="cos t") # axにcos関数も登録

# 以下同様。ただし，ax1, ax2は使わず，axのみを使う
```

Q. 上の例を修正して，1つのグラフに以下の2つの曲線を描きなさい。

グラフ1
- $\displaystyle y=1-\frac{x^2}{2!}$を描画
- 関数`plot()`の`引数label`には`$1-x^2/2!$`を指定。 (`1-x^2/2`を`$`で囲む。グラフ上の表示には\$は無くなるはず )
- 横軸ラベルは`$x$`
- 縦軸ラベルは`$y$`
- タイトルは`Fourier series of cos(x)`
- $-3\le x\le3$,  $-1\le y\le1$の範囲で描画

グラフ2
- $\displaystyle y=1-\frac{x^2}{2!}+\frac{x^4}{4!}$を描画
- 関数`plot()`の`引数label`には`$1-x^2/2!+x^4/4!$`を指定。
- 軸ラベルは指定なし
- 描画範囲指定もなし

共通の注意
- 横軸データはNumpyを使い，$[-3,3]$の区間を0.1ごとに生成したデータを使うこと(端点(-3.0と3.0)も含むこと)
- 格子線も描くこと
- fontsizeは見やすい大きさに調整すること
- `ax`の名称は変えないこと

コメント
- ラベルの文字を`$`で囲むのは，数式を適切に表示するための書式(LaTeX形式)

<!-- grade_id: cell-d64704a41c62a89b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-mpl-label, points: 2, grade: true, solution: false -->
```python
# 評価用セル(軸ラベル確認)

```

<!-- grade_id: E-mpl-data, points: 2, grade: true, solution: false -->
```python
# 評価用セルその2(描画関数確認)
```

<!-- grade_id: E-mpl-title, points: 2, grade: true, solution: false -->
```python
# 評価用セルその3(各曲線ラベル確認)
pc = LinePlotChecker(ax)
pc.assert_title_equal("Fourier series of cos(x)")
pc.assert_labels_equal(['$1-x^2/2!$', '$1-x^2/2!+x^4/4!$'])
```

<div class="alert alert-success">

# ヒートマップ

</div>

```python
%matplotlib inline
import numpy as np
from matplotlib import pyplot as plt
# ここまでは，すでに実行ずみなら不要

data = [[1,2,3,4],
        [2,3,4,5],
        [3,4,5,6],
        [4,5,6,7]]

fig=plt.figure(figsize=(6, 5)) # グラフを生成
ax=fig.subplots()             # 座標軸を1つ生成してaxに代入
# ax.invert_yaxis() # y軸を上下ひっくり返したいとき

map=ax.pcolor(data,cmap="seismic") # 数値データからカラーデータを作成 cmapの値により配色変化
fig.colorbar(map, ax=ax) # カラーデータと軸データを使ってグラフデータ生成

plt.show() # 表示(生成されているグラフデータを画面表示)
```
- cmap(カラーマップ)の一覧は[こちら](http://pynote.hatenablog.com/entry/seaborn-color-map)にあるので，いろいろと変えてみると楽しい

## 軸目盛の間隔
x軸やy軸の目盛(tick)の間隔を変更したいときは，リストで指定する。以下は，0から9までの整数値を軸名に指定する場合の例。
```python
ax.set_xticks([range(0,11,1)])
ax.set_yticks([range(0,11,1)])
```

注) リストではなく，numpy配列でも良い

Q. 上記ヒートマップのプログラムを次のセルに入力しなさい。また，`range()`を用いて0から4の配列を生成し，それを変数`tick`に代入しなさい。
 この変数`tick`を利用し，先述のカラーマップのx軸およびy軸の目盛を1刻みになるよう指定しなさい。
 
 注) 変数`ax`の名前は変えないこと

<!-- grade_id: cell-e4f9dc1a5f670aca, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-mpl-heatmap, points: 2, grade: true, solution: false -->
```python
#  評価用セル
```