<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">

# Pythonの基本(basic)
    
</div>

## Pythonのプログラムの作成と実行

方法1: C言語のように，テキスト形式でプログラムを書いて実行する。

```bash
$ vi program.py # プログラム作成
$ python program.py # プログラム実行
```

方法2: このjupyter notebook等で，インタラクティブに実行する。本演習ではこちらを利用。

<div class="alert alert-success">

# Pythonの基本
    
</div>

## 命令文の書き方 p.44

- C言語の`;`のような**行末記号はない**
- 複数のstatement (命令文)を1行に書くときには `;`で区切る
- statement の途中で改行するときには `\`で，次行に続くことを明示する
```python
a=1
b=2; c=a+b; d=a+b+c
e=a+b \
   +c+d
``` 
注) pythonにはC言語のような**行末記号がない**ので，命令文が複数行にわたるときには，**途中で改行することを`\`で明示**する必要がある

```python
# このセルで試してみよう
```

## コメント p.46
### 1行コメント

`#`から行末までは実行時に無視される。

```python
# shift-Enterを押して実行してみよう
x=3    #  x=3は実行されるが，#以降のコメント部分は無視される
# x=4     この行は実行されない
print(x)
```

### 複数行コメント
複数行にわたったコメントを書くときは`"""`，もしくは，`'''`で挟む。

```python
'''
ここは実行時に無視される
'''
x=5 # ここは実行される

"""
ここも実行時に無視される
x=2
"""
print(x)
```

ただし，以下のようにコメント部分の後に命令文がない場合は，コメント内の文が出力セルに表示される。プログラムの命令文には影響しない。

```python
x=5
"""
ここはコメント文だが，下に命令文が無い場合は，実行時に出力セルに表示される...
"""
```

## print文, p.49
関数`print()`は，`%s`等で型指定しなくても，変数に合わせて適当に表示してくれる。

```python
x=3
print(x)
```

複数の変数を表示する方法はいろいろある。

```python
x=3
y=2.0
print(x,y,sep=",") # 区切り文字を指定できる
print("x=", x, " y=", y)
```

表示のための型指定をすることもできる。

```python
print("x=%d, y=%e" % (x,y))
```

## 基本演算
- Cとおおむね同じ (演算子一覧は p.54, 71)
- ただし，`++`, `--`は使えない。`+=`, `-=`は使える

```python
x=3
x+=3
print(x)
y=x%4
print(y)
```

## 関数の定義方法, p.232

- `def`で始まる行で関数名と引数を定義する
- **次行から関数の終わりまではインデント**して命令文を書く

書式
```python
def function_name(args):
    命令文       #ここから関数の終わりまでインデントをする
    return(..)
```

合計を求める関数と，その実行例

```python
def sum(x,y):
    return(x+y)

print(sum(2,3))
```

Q. 名前`name`, 体重`m`, 身長`h`を引数として受け取り，「◯◯の身長は◯cmで体重は◯kgです」と表示する関数`mybody(name,m,h)`を定義しなさい。


**実行例**
- 命令文
```
mybody("Taro",50,170)
```
- 実行結果
```
Taroの身長は170cmで体重は50kgです
```

<!-- grade_id: A-basic-mybody, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
def mybody(name,m,h):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-basic-mybody, points: 1, grade: true, solution: false -->
```python
# 評価用セル
with capture_stdout() as capture:
    mybody("taro",50,170)
res=capture.result
assert_regex(res.replace(" ","").strip(),re.compile("taroの身長は170cmで体重は50kgです"))

with capture_stdout() as capture:
    mybody("hanako",h=160,m=40)
res=capture.result
assert_regex(res.replace(" ","").strip(),re.compile("hanakoの身長は160cmで体重は40kgです"))
```

## 変数の型(p. 50)
- 変数の型宣言は不要(型は自動推定)
- type(x) : xの型名を返す
```python
x=2
y=2.0
z=3+2j
l=[3,5,2]
b=True
s="hello"
```
例えば，上記の変数`x`の型は以下を実行すると確認できる。

```python
x=2
print(type(x))
```

実行結果より，`x`の型は`int`であることがわかる。

Q. 上記の変数`z`, 変数`b`, 変数`s`の型を確認しなさい。その結果をそれぞれ変数`type_z`, `type_b`, `type_s`に代入しなさい。


**解答例:** 以下は上記の変数`x`の型を答える場合。
```python
type_x="int"
```

```python
# このセルで確認しなさい
```

<!-- grade_id: A-basic-type, points: 0, grade: false, solution: true -->
```python
# こちらに答を入力
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-basic-type, points: 0.5, grade: true, solution: false -->
```python
# 評価用セルその1
```

<!-- grade_id: E-basic-type2, points: 0.5, grade: true, solution: false -->
```python
# 評価用セルその2
```

<!-- grade_id: E-basic-type3, points: 0.5, grade: true, solution: false -->
```python
# 評価用セルその3
```

## 型の判定

`type(x)`の返り値によって，`x`が`bool`か否かは以下のように判定できる。
```python
x=3
print(type(x) is bool)
print(type(x) is not bool)
```

Q. 引数`x`の型が`int`ならば`True`を， それ以外は`False`を返す関数`is_int()`を作りなさい。

**注意**
- 答を表示するのでは無く，返り値を設定すること
- 返り値は，文字列にはしないこと("True"や"False"とダブルクォーテーションで囲むと文字列になる)。

<!-- grade_id: A-basic-is_int, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
def is_int(x):
   # YOUR CODE HERE
   raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-basic-is_int, points: 1, grade: true, solution: false -->
```python
# 評価用セル
eq_(is_int(3),True)
eq_(is_int(-1),True)
eq_(is_int("cat"),False)
# 非公開の評価プログラムもあるので，(この問に限らす)いろいろな場合の動作確認をしておくこと
```

<div class="alert alert-success">

# モジュール

</div>

C言語のライブラリにあたるもの

## mathモジュール p.78
- 関数例(一覧p.80-81)
    - math.sin()
    - math.cos()
    - math.log()
- 定数
    - math.pi
    - math.e
    
詳細情報を知りたいときには以下のように`help()`を使う
```python
help(math)
help(math.sin)
```

mathモジュールの読み込みと関数実行
```python
import math as ma
x=ma.pi/2
ma.sin(x)
```

**注意:**
jupyter notebook上で一度`import`したモジュールは，他のセルでも使える。複数のセルで同じモジュールを何度も`import`する必要はない。

```python
# このセルで試してみよう
```

## randomモジュール p.82
- 関数例
    - `random.randint(min,max)`   # randomモジュールのrandint()関数
    - `random.random()`

```python
import random as ra # モジュールの読み込み
x=ra.random()
print(x)
```

Q. `random`モジュールの関数`random.randint(0,10)`および`random.random()`により生成される乱数の最大値と最小値をそれぞれ調べて以下に書きなさい。

解答例)
```python
randint_min=3     # 関数random.randint(0,10)の出力の最小値
randint_max=4     # 関数random.randint(0,10)の出力の最大値
random_min=5      # 関数random.random()の出力の最小値
random_max=6     # 関数random.random()の出力の最大値(上限が6.0未満の場合は6と答える)
```

<!-- grade_id: A-basic-rand_range, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-basic-randint_range, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
```

<!-- grade_id: E-basic-random_range, points: 0.5, grade: true, solution: false -->
```python
# 評価用セルその2
```

## モジュールのいろいろな読み込み方法, p.81
```python
import <module>
import <module> as <alias>
from <module> import <function>
```
例)
```python
import math
print(math.cos(0))

import math as ma
print(ma.cos(0))

from math import cos
print(cos(0))
```

Q. `random`モジュールの関数random.random()を，単に`random()`と書くだけで呼び出せるように，`random`モジュールを`import`しなさい。

<!-- grade_id: A-basic-import_random, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルでrandom()の動作確認をしよう
```

<!-- grade_id: E-basic-import_random, points: 0.5, grade: true, solution: false -->
```python
# 評価用セル
assert random()
```

<div class="alert alert-success">

# ループ・条件分岐
    
</div>   

## if文 p.96
```python
x=1
if x == 0:
    print(x, "is zero")
elif x > 0:
    print(x, "is positive")
else:
    print(x, "is negative")
```
注) pythonはC言語のように`{}`を使わない。インデント(空白4文字)で構文単位を表現

Q. 引数`x`が偶数なら0を， 奇数なら1を，それ以外の場合は2を返す関数 `parity()`を作りなさい。
- 色々な引数を想定して，どのような引数に偶奇の判定ができない場合はどういう時かよく考えること。
- 一例として，真理値(True/False)には偶奇判定はできない。計算上は，True/Falseを1/0に対応づけることもあるが，対応づけの規則はプログラムミング言語により異なり，また，必ずしも1対1の関係ではない。例えば，Pythonでは0以外の数値は全てTrueに対応づけられる。

<!-- grade_id: A-basic-parity, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
def parity(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
type(True)
```

<!-- grade_id: E-basic-parity, points: 1, grade: true, solution: false -->
```python
# 評価用セル
eq_(parity(6),0)
eq_(parity(-10),0)
eq_(parity(11),1)
eq_(parity(1.1),2)
eq_(parity(-15),1)
eq_(parity(True),2)
# 他にも偶奇判定できない例は色々あります
# 非公開の評価プログラムもあるので，(この問に限らす)いろいろな場合の動作確認をしておくこと
```

## while文, p.112
```python
i = 0
while i < 10:
    print(i, end=' ')
    i += 1
```

Q. 1から引数で与えた整数`x`までの整数を順に足して，最後にその総和を返す関数`total_int()`をwhile文を使って書きなさい。

<!-- grade_id: A-basic-total_int, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力
def total_int(max):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-basic-total_int, points: 1, grade: true, solution: false -->
```python
# 評価用セル
eq_(total_int(3),6)
eq_(total_int(10),55)
# 非公開の評価プログラムもあるので，色々な場合の検証をすること
```

## for文, p.121
### リストを使う

リストは要素を`[]`内に列挙してつくる。要素の取り出しは，例えば以下のように`for`文で行える。
リストの詳細は今後の課題でも扱う。

```python
for x in [1,5,9]:
    print(x, end=' ')
```

リストは変数に代入することもできる。

```python
numlist=[1,5,9]
for i in numlist:
    print(i, end=' ')
```

Q. 引数として受け取った文字列リスト`s`(例: `fruit=["apple", "orange", "banana"]`)の各要素を順に表示する関数`show_list()`を書きなさい。
すなわち，
```python
fruits=["apple", "orange", "banana"]
show_list(fruits)
```
を実行すると，以下のような1行に出力されるとする。(区切り文字は空白にしなさい)
```python
apple orange banana
```

<!-- grade_id: A-basic-showlist, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
def show_list(s):
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
fruits=["apple", "orange", "banana"]
show_list(fruits)
```

<!-- grade_id: E-basic-showlist, points: 1, grade: true, solution: false -->
```python
# 評価用セル
with capture_stdout() as capture:
        show_list(["apple", "orange", "banana", "kiwi"])

assert str.strip(capture.result) =="apple orange banana kiwi"
```

### range()関数を使う

**`range(end)`** :0以上end**未満**の数列を生成する

```python
for i in range(5):
    print(i, end=' ')
```

**`range(begin, end)`**: begin以上end**未満**の数列を生成する。

```python
for i in range(1,5):
    print(i,end=' ')
```

**`range(begin, end, step)`**: begin以上end**未満**の数列を生成する。stepは要素間の間隔

```python
for i in range(3,8,2):
    print(i, end=' ')
```

Q. 正の引数`n`に対して， `n n-1 n-2 ... 1`を順に`print()`で表示する関数`decrement(n)`を`range()`を用いて作りなさい。

実行例
```python
decrement(5)
```
出力表示
```python
5 4 3 2 1
```

<!-- grade_id: A-basic-decrement, points: 0, grade: false, solution: true -->
```python
def decrement(n):
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-basic-decrement, points: 1, grade: true, solution: false -->
```python
# 評価用セル
with capture_stdout() as capture:
        decrement(5)
assert str.strip(capture.result) =="5 4 3 2 1"
with capture_stdout() as capture:
        decrement(10)
assert str.strip(capture.result) =="10 9 8 7 6 5 4 3 2 1"
```

### for文の一般的な表現
```python
for 変数 in イテラブル:
    実行文
```
- **イテラブル**: 順に一つずつ要素を取り出せるオブジェクト
    - ex) リスト, range()でつくるシーケンスなど
- for と inの 間の変数に，イテラブルの要素が順に代入される

### 2重ループの例
```python
for i in range(3):
    for j in range(4):
        print("(%d, %d)" % (i,j))
```

Q. 数字リスト(例: `states=[1,2,3,4,5]`)と文字リスト(例: `actions=["right","left"]`)を引数とし，各要素のすべての組み合わせについて，以下のように表示をするfor文を作りなさい。
- 変数`states`に代入するリスト(数列)はrange()を使って生成しなさい。

実行例
```python
states=range(1,4)
actions=["right","left"]
combination(states,actions)
```
実行結果
```python
(1,right)(1,left) #<=数字が繰り上がる毎に改行
(2,right)(2,left)
(3,right)(3,left)
```

<!-- grade_id: A-basic-combination, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
def combination(states,actions):
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
states=range(1,4)
actions=["right","left"]
combination(states,actions)
```

<!-- grade_id: E-basic-combination, points: 1, grade: true, solution: false -->
```python
# 評価用セル

with capture_stdout() as capture:
        combination(range(1,2),["top","down"])
res=capture.result
assert res.replace(" ","").strip() =="(1,top)(1,down)"

with capture_stdout() as capture:
        combination(range(1,4),["right","left"])
res=capture.result
assert res.replace(" ","").strip() =="(1,right)(1,left)\n(2,right)(2,left)\n(3,right)(3,left)"
```