<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">

# ユーザ定義関数(function), p.232
    
</div>

## 関数の定義

```python
def function_name(args):
    命令文
    return(..)
```
以下は簡単な関数例

```python
def print_name(family,first): # family: 姓,  first: 名
    print("My name is %s %s." % (first, family))
```

C言語と同様に引数を設定して呼び出せば実行できる。

```python
print_name("Yamaguchi","Daisuki")
```

引数名をつけて，引数の値を指定することもできる。

```python
print_name(family="Yamaguchi",first="Daisuki")
```

引数名をつければ，引数の順番はランダムでも良い

```python
print_name(first="Daisuki",family="Yamaguchi")
```

## 引数のデフォルト値, p.243

引数にデフォルト値を設定することもできる。

```python
def print_name(family="Momo",first="Taro"): # family: 姓,  first: 名
    print("My name is %s %s." % (first, family))
```

デフォルト値の設定されている引数は，呼び出し時に省略可能。以下では`family`の方のみ省略している。

```python
print_name(first="Daisuki")
```

## if-elseを1行で書く p.110

**1行if文の書式** : `返り値 if 条件式 else 返り値`

例えば，絶対値を求める関数`absval`を`def`を使って書くと以下のようになる。
```python
def absval(x)
    if x>0:
        absval=x
    else:
        absval=-x
    return(absval)
```

この関数を1行にまとめて書くと，以下のようにとても簡潔になる。

```python
def absval(x):
    return(x if x>0 else -x)
```

```python
# 実行例
print(absval(-1))
```

Q. 確率`epsilon`でTrueを返し，確率`1.0-epsilon`でFalseを返す関数`eps_random(epsilon=0.1)`を定義しなさい。

```python
import random as ra
def eps_random(epsilon=0.1):
    # ra.random()で乱数を発生し，epsilonよりも小さければTrueを，それ以外はFalseを返す処理を書く。
```

<!-- grade_id: A-erandom, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
import random as ra
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-erandom, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## 複数の返り値をもつ関数

Pythonの関数では，返り値を複数指定できる。

```python
def interval(mean,sd): # 平均mean, 標準偏差sdから，サンプルの約68%が含まれる区間を出力
    return mean-sd, mean+sd
```

複数の返り値を，それぞれ別の変数で受け取ることができる。

```python
low,high=interval(70,10)
print(low,high)
```

複数の返り値を，一つの変数でまとめて受け取ることもできる。

```python
ret=interval(70,10)
print(ret[0], ret[1])
print(type(ret))      # 返り値はタプル(要素の書き換えを出来ないリスト。タプルについては後述)
# ret[0]=3            # 返り値の要素を変更するとエラーが出る。(タプルの要素は書き換え不可。後述)
```

## immutable と mutable

関数の引数(仮引数)の値を関数内で変更した場合，引数が**Immutable(変更不可)**な型ならば呼び出した側の引数(実引数)は変化しない(値渡し)。**Mutable(変更可能)**ならば，実引数も変化する(参照渡し)。

- **immutable**: int, float, str, tuple 等
- **mutable**: list, dict, set 等(これらについては後述)
   
immutableな引数については， 返り値を使って書き換えようにする。

Q. 下記のプログラムの実行後，変数stateの値は変化しているだろうか。言い換えると，関数`new_state()`に渡した引数`s`はimmutableだろうか，それともmutableだろうか? 解答欄には以下のように記載しなさい。

```python
def new_state(s,a):
    if a == "left":
        s-=1
    else:
        s+=1
    return(s)

state=5
new_state(state,"right")
```
解答例(以下のいずれかを解答する)
```python
answer="I"  # immutableの場合
answer="M"  # mutableの場合
```

```python
# このセルで試してみよう
```

<!-- grade_id: cell-8eaff5a058528024, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-883b2564dfa2bc75, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# よくあるトラブル

</div>

標準関数名と同じ名前の関数や変数を作ってしまうと，標準関数にアクセスできなくなるので要注意。

関数名の頭に?をつけて，例えば
```python
?sum
```
として実行すると，指定した関数の使い方が表示される。動作が怪しいときには，これを使って参照されている関数の確認できる。

定義した変数や関数を消去したいときには`del()`を使う

Q(1) Pythonの標準関数に最大値を返す関数`max()`がある。
1. `max(3,5)`の返り値が5であることを，次のセルで実行して確認しなさい。

```python

```

Q(2) `max=3`を実行した後に `max(3,5)`を実行するとどのようなエラーがでるか，次セルで実行して確認しなさい。
```python
max=3
print(max(3,5))
```

<!-- grade_id: cell-eae1aa59539e2a56, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Q(3) 前のセルの'max=3'の次の行に，変数`max`を削除して標準関数`max()`を再び使えるようにする命令を追加しなさい。その結果，エラーが出なくなることを確認しなさい。
```python
max=3
ここに何か命令文を入れる
print(max(3,5)) # エラーが出なくなることを確認する
```

<div class="alert alert-success">

# 練習問題

</div>

一列に並んだ`N`個の部屋`(s=1,2...,N)`があり，その中にねずみ(エージェント)がいる。
- いずれかの部屋の中にいるエージェントは行動`a`として"left"か"right"を選ぶことができ，それぞれ1だけ部屋番号が小さい，もしくは大きいほうの隣の部屋に移動する。
- 部屋`1`で"left"，もしくは部屋`N`で"right"を選択すると，エージェントは部屋の外に出てエピソード終了となる。
- 部屋`N`で行動"right"を選択したときのみ，エージェントは報酬(reward)として10を得ることができる。それ以外の場合は報酬rewardは0である。

このような環境について，以下の一連の関数を作りなさい。

Q1) 上述の環境における状態遷移を表現する，以下のような関数`step(s,a)`を作りなさい。

- 概要
    - 引数として受け取った現在の状態`s`と選択した行動`a`により状態遷移をした後の新しい状態`s`および，状態遷移によって獲得する報酬を返す。
- 引数
    - `s`: 現在の状態(部屋番号)
    - `a`: 選択した行動
- 返り値
    - `s`: 新しい状態。ただし，エピソード終了時の状態値は，部屋`1`から左に出た場合は`0`，部屋`N`から右に出た場合は`N+1`とする。
    - `r`: 報酬
    - `terminal`: エピソード終了フラグ。エピソード終了時には`True`, それ以外には`False`

<!-- grade_id: cell-ff9f3173d298c68c, points: 0, grade: false, solution: true -->
```python
N=5 # Nが他の値でも正常動作するようにすること
def step(s, a):
    # YOUR CODE HERE
    raise NotImplementedError()
    return s, r, terminal
```

```python
# このセルで動作確認をしよう(Nが5以外でも動作すること)
```

<!-- grade_id: cell-2b2c4ebd1626cf58, points: 2, grade: true, solution: false -->
```python
# 評価用セル
N=20
s, r, terminal=step(5,"right")
assert_equal(s,6)
assert_equal(r,0)
assert_equal(terminal,False)
s, r, terminal=step(20,"right")
assert_equal(s,21)
assert_equal(r,10)
assert_equal(terminal,True)
```

<!-- grade_id: cell-80314bf21a4f7342, points: 2, grade: true, solution: false -->
```python
# 評価用セルその2
```

Q2) エージェントがどの部屋にいるかを以下のように表示する関数`render()`を作りなさい。

- 引数
    - s: エージェントの状態(部屋番号)
    - r: 報酬(デフォルト値0)
- 返り値
    - なし
- 実行内容
    - 以下のように，部屋番号と，引数で与えられたエージェントの位置(@マークで表す)を表示
    - 右外には，引数で与えられた報酬値rを表示

**実行例**) N=5の場合
```python
N=5
render(2) # state=2の場合
render(s=6,r=10) # state=6 (エピソード終了時)
```
表示結果
```python
  |1|@|3|4|5| r=0
  |1|2|3|4|5|@ r=10
```

**実行例2**) N=3の場合
```python
N=3
render(s=4,r=10) #  state=4 (エピソード終了時)
render(0) # state=0 (エピソード終了時)
```
表示結果
```python
  |1|2|3|@ r=10
 @|1|2|3| r=0
```

<!-- grade_id: cell-81f39747feef58b1, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-2af6ef6de76b4c00, points: 2, grade: true, solution: false -->
```python
# 評価用セル
N=5
with capture_stdout() as capture:
        render(2)
res=capture.result
assert res.replace(" ","").strip() =="|1|@|3|4|5|r=0"
with capture_stdout() as capture:
        render(6,10)
res=capture.result
assert res.replace(" ","").strip() =="|1|2|3|4|5|@r=10"
```

<!-- grade_id: cell-61004cccd78c2337, points: 2, grade: true, solution: false -->
```python
# 評価用セルその2
```

Q3) Q1, Q2でつくった関数`step()`, `render()`のデバッグのため，以下のような関数`trial(a)`を作りなさい。

- 概要
    - 関数`step()`呼び出しながら，エピソード終了までの状態遷移の様子を`render()`を用いて表示する。エージェントがとる行動は，常に引数`a`で指定したものとする。
- 引数
    - `a`: エージェントがとる行動`a`。**デフォルト値は"right"**
- 処理内容
    - 引数`a="right"`の場合は状態`state=1`からエピソードを開始し，
    ，`a="left"`の場合は状態`state=N`からエピソードを開始する。
    - 状態遷移には，Q1で作った関数`step()`を用いて新しい状態や報酬等の情報を受けとり，各ステップにおけるエージェントの位置をQ2で作った`render(s,r)`を用いて表示する。
    - エージェントのとる行動は，常に引数`a`で受け取った行動とする。
    - 関数の終了は，エピソード終了時。
- 返り値
    - なし

**実行例**)
```python
N=5
trial()
```
表示結果
```python
  |@|2|3|4|5| r= 0
  |1|@|3|4|5| r= 0
  |1|2|@|4|5| r= 0
  |1|2|3|@|5| r= 0
  |1|2|3|4|@| r= 0
  |1|2|3|4|5|@ r= 10
```

実行例2)
```python
N=5
trial("left")
```
表示結果
```python
  |1|2|3|4|@| r= 0
  |1|2|3|@|5| r= 0
  |1|2|@|4|5| r= 0
  |1|@|3|4|5| r= 0
  |@|2|3|4|5| r= 0
 @|1|2|3|4|5| r= 0
```

<!-- grade_id: cell-97a7e577edb518a5, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: cell-f34f182ab59d4800, points: 2, grade: true, solution: false -->
```python
# 評価用セル
N=3

with capture_stdout() as capture:
        trial()
res=capture.result
assert res.replace(" ","").strip() =="|@|2|3|r=0\n|1|@|3|r=0\n|1|2|@|r=0\n|1|2|3|@r=10"

with capture_stdout() as capture:
        trial("right")
res=capture.result
assert res.replace(" ","").strip() =="|@|2|3|r=0\n|1|@|3|r=0\n|1|2|@|r=0\n|1|2|3|@r=10"

with capture_stdout() as capture:
        trial("left")
res=capture.result
assert res.replace(" ","").strip() =="|1|2|@|r=0\n|1|@|3|r=0\n|@|2|3|r=0\n@|1|2|3|r=0"
```

<!-- grade_id: cell-9c1363e2e4ec0ca9, points: 2, grade: true, solution: false -->
```python
# 評価用セルその2
```