<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">

# Numpyモジュール, p.344

</div>

## Numpy とは

- **N次元配列(N-dimensional array, ndarray)**を定義し， **ベクトルや行列をはじめとする数値計算を高速に行うための**モジュール。
- NumPyモジュールで生成する<span class="mark">配列の型名は`numpy.ndarray`</span>

参考資料
- 以下のNumPyの説明は，[Python Data Science Handbook: "The Basics of NumPy Arrays"](https://github.com/jakevdp/PythonDataScienceHandbook/blob/master/notebooks/02.02-The-Basics-Of-NumPy-Arrays.ipynb)をいろいろと参考にしています。
- NumPyのURLは [www.numpy.org](http://www.numpy.org/)


## モジュールの読み込み

NumPyモジュールは，慣例的に`np`の名称で利用される。

```python
import numpy as np

print(np.__version__)
```

## NumPyのマニュアル

以下で`NumPy`のマニュアルを表示できる。

```python
np?
```

## NumPyの関数とメソッド

- 配列`ndarray`の**生成**や，**複数の配列に対する操作**(結合等)には, NumPyの**関数**を使う
- すで生成されている`ndarray`オブジェクトの**加工**には，`ndarray`に対する**メソッド**を使う
- 以下では，NumPyの関数とメソッドの説明の際に，両者を以下のように書いて区別する
    - **関数**: `numpy.array()`等と，モジュール名`numpy`もしくは略称`np`をつけて説明
    - **メソッド**: `ndarray.reshape()`等と`ndarray`をつけて説明

<div class="alert alert-success">

# 配列`ndarray`の生成, p.344

</div>

## 関数`numpy.array()`

リスト(list)を`numpy.array()`で`ndarray`に変換する。

```python
x = np.array([1, 2, 3])
print(x)
```

```python
l = [1, 3, 5]
y = np.array(l)
print(y)
```

```python
np.array([[0, 1], [2, 3]])
```

タプル(tuple)を`numpy.ndarray`に変換する

```python
t = (2, 4, 6)
z = np.array(t)
print(z)
```

ここまでに生成した変数の型が`numpy.ndarray`であることを確認する。

```python
print(x, type(x))
print(y, type(y))
print(z, type(z))
```

## 配列`ndarray`に関する情報

配列`ndarray`の情報(次元，形状，要素数)を取得するには，`ndarray`の属性(`ndim`, `shape`, `size`)にアクセスすれば良い。

```python
x = np.array([1, 2, 3])
print(x)
print("ndim:", x.ndim)  # 次元
print("shape:", x.shape)  # 形状
print("size: ", x.size)  # 要素数
```

上の`x`は1次元配列なのでndimは1。shapeは`(3,)`と表示され，2次元目の情報は表示されていない。

以下は2次元配列の例。ndimが2, shapeは`(3,2)`と，2次元配列となっていることがわかる。

```python
x = np.array([[0, 1], [2, 3], [4, 5]])
print(x)
print("ndim:", x.ndim)  # 次元
print("shape:", x.shape)  # 形状
print("size: ", x.size)  # 要素数
```

## 要素の型

リスト(list)の要素はいろいろな型を混在させることができる。一方，Numpyの配列の各要素は同じ型でないといけない。

要素に`int` と `float` を混在させて定義した場合，要素の方は全て`float`に変換される。

```python
y = np.array([1, 2.1, 3])
print(y.dtype)
```

Numpyオブジェクトの要素の型は引数`dtype`を用いて陽に指定することもできる。

```python
x = np.array([1, 2.1, 3], dtype="int")
print("x=", x)
print(x.dtype)  # 要素の型の確認
```

## `ndarray`のいろいろな生成方法, p.372

要素数10の配列。要素値の初期化はしない。

```python
np.empty(10)
```

全ての要素値が0, 要素数10の配列(`numpy.zeros()`を使う)

```python
np.zeros(10)
```

全ての要素値が0の2x3の二次元配列

```python
np.zeros((2, 3))
```

全ての要素値が1, 要素数10の配列。以下では要素の型を`int`に指定。

```python
np.ones(10, dtype=int)
```

全ての要素値が1の 3x4の2次元配列

```python
np.ones((3, 4))
```

全ての要素値が`np.pi`である要素数5の配列

```python
np.full(5, np.pi)
```

**等差級数の生成(間隔(step)を指定)**

- `np.arange(end)` : \[0,end) の区間の整数値生成
- `np.arange(start,end)`: \[start,end) の区間の整整数値生成
- `np.arange(start,end,step)` : \[start,end) の区間内でstepごとの数値を生成


以下は，各要素の値が1から6の整数である要素数6の配列を浮動小数点に変換している。

```python
np.arange(1, 7, dtype=float)
```

各要素の値が0以上2**未満**の範囲で0.2間隔で数列を生成

```python
np.arange(0, 2, 0.2)
```

**等差級数の生成(分割数を指定)**

`numpy.linspace(start,end,num)`: [start, stop]の区間を端点を含む`num`個の境界値で分割し，その境界値を返す。

0以上2**以下**の区間を等間隔に区切る11個の値を要素とする配列(`numpy.linspace()`を使う)

```python
np.linspace(0, 2, 11)
```

3x3の恒等行列

```python
x = np.eye(3)
y = np.identity(3)
print(x)
print(y)
```

Q. 0から63までの整数値を要素とする`ndarray`を作り，変数`arr`に代入しなさい。

<!-- grade_id: cell-0b24f7d6ffb76616, points: 0, grade: false, solution: true -->
```python
# 解答欄 (arr=...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-nparange, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 全ての要素値が**整数値**の0である，要素数20の`ndarray`を作り，変数`arr20`に代入しなさい。

<!-- grade_id: cell-aae027fc195a8d9d, points: 0, grade: false, solution: true -->
```python
# 解答欄 (arr=...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-npzeros, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. $[-5,5]$の区間を等間隔に0.5ごとに区切った場合に境界値になる値を要素とする`ndarray`を作り，変数`arr`に代入しなさい。
上限値と下限値も要素に含めること

<!-- grade_id: cell-804a7953d7655ff4, points: 0, grade: false, solution: true -->
```python
# 解答欄 (arr=...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-split05, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. $[-2,2]$の区間内を等間隔に7つの区画に分割する場合に境界値となる値を要素とする`ndarray`をつくり，変数`arr`に代入しなさい。
上限値と下限値も要素に含めること

<!-- grade_id: cell-71d0247a71201e1e, points: 0, grade: false, solution: true -->
```python
# 解答欄 (arr=...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-split8, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## 乱数による配列の生成, p.377

要素値が0から1の一様乱数である，2x3の2次元配列を作る。

```python
np.random.rand(2, 3)
```

0以上5**以下**の整数乱数を要素とする, 2x2x2の3次元配列を作る。

```python
x = np.random.randint(0, 6, (2, 2, 2))
y = np.random.randint(6, size=(2, 2, 2))
print("x=", x)
print("y=", y)
```

平均0, 標準偏差1の正規分布乱数を要素とする，要素数20の配列を作る。

```python
L = np.random.normal(0, 1, 20)
print(L)
print("mean:", np.mean(L))  # Lの平均値
print("SD:  ", np.std(L))  # Lの標準偏差
```

以下でも，同様な配列を作れる。

```python
L = np.random.randn(20)
print(L)
print("mean:", np.mean(L))  # Lの平均値
print("SD:  ", np.std(L))  # Lの標準偏差
```

前セルと同様に平均0, 標準偏差1の正規分布乱数を要素とする，要素数1000の配列を作り，ヒストグラムに示す方法は以下のようになる。

```python
import matplotlib.pyplot as plt

L = np.random.randn(1000)

fig = plt.figure()
ax = fig.add_subplot(111)
ax.hist(L, bins=20)  # binsがビン数(横軸を分割した区間の数)
fig.show
```

Q. $[-10,10]$の範囲の整数乱数を要素とする, 要素数300の配列を作り，変数`rand10`に代入しなさい。(生成された乱数の範囲が適切か確認すること)

<!-- grade_id: cell-3d6175d34c6d9dbf, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-rand10, points: 1, grade: true, solution: false -->
```python
# 評価用セル
eq_(max(rand10), 10)
eq_(min(rand10), -10)
```

Q. 平均値を5.0, 標準偏差を1.0の正規分布乱数を要素とする, 要素数1000の配列を作り，変数`norm5`に代入しなさい。

<!-- grade_id: cell-80876db664f57cb1, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-norm5, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# 配列の参照, p.353

</div>

## 1次元配列の参照

NumPy配列要素の値を参照する方法はリスト(`list`)と同様。すなわち，インデックスに`[begin:end:step]`と範囲指定することで行える。

以下で定義する配列`x`を例に紹介する。

```python
x = np.arange(1, 10)
```

`x`の要素一覧

```python
print(x)
```

`x`の，一番はじめの要素

```python
x[0]
```

`x`の，最後から2番目の要素

```python
x[-2]
```

はじめの3つの要素

```python
x[:3]
```

最後の3つの要素

```python
x[-3:]
```

はじめの要素から，一つおきに要素値を表示

```python
x[::2]
```

要素を逆順に表示 (インデックス指定のステップ値に負の値を使う)

```python
x[::-1]
```

## ２次元配列の参照

`ndarray`の2次元配列では，行および列それぞれに対する範囲を`[行に対する範囲,列に対する範囲]`と指定することで参照・スライスできる。

要素一覧

```python
N = np.array([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]])
print(N)
```

初めの行のみ

```python
print(N[0, :])
print(N[0])
```

最後の列のみ

```python
N[:, -1]
```

初めの2行かつ，はじめの3列以内の要素

```python
N[:2, :3]
```

奇数番目の列と全ての行要素できる行列

```python
N[:, ::2]
```

`N`の右下部分の要素でできる2x2部分行列(インデックスに負の値を使う)

```python
N[-2:, -2:]
```

行を逆順にする。(列順は変えない)

```python
N[
    ::-1,
]
```

Q. 以下の配列`I`から奇数行と偶数列のみを抜き出し，変数`Q`に代入しなさい。

<!-- grade_id: cell-23c41767a15c98f3, points: 0, grade: false, solution: true -->
```python
I = np.eye(10)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-slice-odd-even, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 以下の配列`P`から, 右上部分の2行3列を抜き出した配列`Q`を作りなさい。

<!-- grade_id: cell-ecc2a81dc525287a, points: 0, grade: false, solution: true -->
```python
P = np.arange(25).reshape(5, 5)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-slice23, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# 配列のコピー

</div>

Q. 以下を実行すると，`p`と`q`の要素値は最終的にはどうなるだろうか。
```python
x2 = np.array([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]])
x3 = x2
x2[0, 0] = 100
x3[-1, -1] = 200
```
解答欄に以下のように答えなさい。
```python
ans=True # x2とx3が等しいとき
ans=False # x2とx3が等しくないとき
```

```python
# このセルで試してみよう
```

<!-- grade_id: cell-bd9ecacbd3adaf92, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-sustitution, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. 以下を実行すると，`x3`の要素値はどうなるだろうか。
```python
x2 = np.array([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]])
x3 = x2.copy()
x2[0, 0] = 100
```
解答欄に以下のように答えなさい。
```python
ans = True # x3がx2と等しいとき
ans = False # x3がx2と等しくないとき
```

```python
# このセルで試してみよう
```

<!-- grade_id: cell-e63dccd75e58b278, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-copy, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

<div class="alert alert-success">

# 配列`ndarray`の変形と要素の追加

</div>

## `ndarray`の変形, p.348

**1次元配列の多次元化**: `ndarray.reshape()`

`ndarray.reshape()`メソッドにより，例えば要素数9の1次元配列を3x3の2次元配列に変形できる。ただし，`ndarray.reshape()`メソッドは**要素を格納するアドレスは変更しない**で，配列の属性`shape`のみを変える。

```python
x = np.arange(9)
print(x)
y = x.reshape(3, 3)
print(y)
```

NumPyの関数`numpy.arange()`で生成した`ndqrrqy`をメソッド`ndarray.reshape()`で変形する場合，以下のようにまとめて書くことができる。

```python
np.arange(9).reshape(3, 3)  # np.arange(9)で生成されたオブジェクトに対してndarray.reshape()を適用
```

`x = np.arange(1, 9)`を列数1の2次元配列に変形

```python
x = np.arange(1, 9).reshape(8, 1)
print(x)
```

上記と同じ操作を，以下の方法で実現することもできる。

```python
x = np.arange(1, 9)
y = x[:, np.newaxis]
print(y)
```

Q. 以下を実行した場合，`x`, `y`の値はどうなるだろうか。
```python
x = np.arange(9)
y = x.reshape(3, 3)
y[0, 0]=10
```
解答欄には，以下のように解答しなさい。
```python
ans = True   # xとyの第一要素(2次元配列の場合は，0行0列目の要素)が等しい場合
ans = False  # xとyの第一要素が異なる場合
```

```python
# このセルで試してみよう
```

<!-- grade_id: cell-e90d079670dd29cc, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-reshape, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## 要素の追加, p.349

`np.append(x, y)`で，配列`x`に対して配列`y`を追加した配列を得ることができる。

```python
x = np.array([1, 2, 3])
y = np.array([8, 9])
x = np.append(x, y)
print(x)
```

配列`L=np.zeros(10)`に，`x=np.ones(10)`を追加することで，以下のような2次元配列を作ることもできる。
```python
[[0. 0. 0. 0. 0. 0. 0. 0. 0. 0.]
 [1. 1. 1. 1. 1. 1. 1. 1. 1. 1.]]
```
この場合，単に`np.append(L,x)`を実行すると，単に要素数20の1次元配列ができてしまう。まず，それぞれを2次元配列に変換する。

Q. `L`, `x`をそれぞれ1行10列の2次元配列に変換しなさい。

<!-- grade_id: cell-f8e106e02be71b6a, points: 0, grade: false, solution: true -->
```python
# 解答欄
L = np.zeros(10)
x = np.ones(10)
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで結果を確認しよう
```

<!-- grade_id: E-reshape1x10, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

ndarray型2次元配列`p`に対して，行数と列数が等しい配列`q`を結合する方法は以下の通り。

1. 行数は変化させず，列数を増やす方向に`q`を結合する場合: `np.append(p,q)`もしくは`np.append(p,q,axis=1)`
1. 列数は変化させず，行数を増やす方向に`q`を結合する場合: `np.append(p,q,axis=0)`


Q2)`np.append()`により，`L`に対して`x`を新しい行として追加し，その結果を`L`に代入しなさい。

<!-- grade_id: cell-c7510446332a94df, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
print(L)
```

<!-- grade_id: E-append_axis0, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

Q. さらに，配列`L`に全ての要素値が2の行を追加しなさい。

<!-- grade_id: cell-b8b17b03a4c841b2, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
print(L)
```

<!-- grade_id: E-append2_axis0, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

ここまでで求めた`L`に対して，全要素の平均と標準偏差を求める方法。

```python
print("mean:", np.mean(L))
print("SD:  ", np.std(L))
```

各列の平均と標準偏差を求める方法。各列に対する演算をするには`axis=0`を指定する。(`axis=0`の場合，行数が増える方向，つまり配列`L[a][b]`に対しては`a`を0から順に変化させるように演算が適用される)

```python
print("mean:", np.mean(L, axis=0))  # axis=0 は列方向に加算
print("sd:  ", np.std(L, axis=0))
```

各行の平均と標準偏差を求める方法。各行に対する演算をするには`axis=1`を指定する。(`axis=1`の場合，列数が増える方向，つまり配列`L[a][b]`に対しては`b`を0から順に変化させるように演算が適用される)

```python
print("mean:", np.mean(L, axis=1))  # axis=0 は行方向に加算
print("sd:  ", np.std(L, axis=1))
```

<div class="alert alert-success">

# 配列`ndarray`の計算

</div>

C言語で配列の演算をするときには，for文による処理が必要だが，NumPyでは，for文を使わないで配列の全要素に対する演算をできるメソッドが用意されている。

**重要:** 配列に対する演算は**for文を使うと演算時間が非常に長くなる**。できるだけ`for`文は使わない方法を用いること。

## 四則演算

ndarray配列の四則演算は，以下のように演算子(`+`,`-`,`*`等)を使う方法とと関数を使う方法がある。

- ndarray**配列同士の四則演算は各要素同士の演算**となる。
- ndarray**配列に対するスカラー値の加減乗除は配列の各要素に対する演算**となる。

演算子|関数
----|--------------
`+` |`np.add()`
`-` |`np.subtract()`
`-` |`np.negative()`
`-` |`np.multiply()`
`/` |`np.divide()`
`//`|`np.floor_divide()`
`%` | `np.mod()`
`**`|`np.power()`
```

以下の演算の結果，要素値がどう変わっているか確認しなさい。

```python
x = np.arange(1, 6)
print(x)
print(x + 2)
print(x * 2)
print(x / 2)
print(x // 2)
print(x % 2)
print(x ** 2)
```

## いろいろな関数

ndarrayに対する演算用にいろいろな関数がある。いずれも，引数にして指定したndarrayの全要素に対して演算が行われる。

```python
np.abs()
np.exp()
np.exp2()
np.power()
np.sin()
np.cos()
np.tan()
np.ln()
np.log()
np.log2()
```
数学定数も色々定義されている。
```python
np.e  # ネイピア数(自然対数の底)
np.pi # 円周率
```

$y=\sin 2\pi x$, $x\in[-1,1]$をmatplotlibでグラフ表示するには以下のように行う。
1. 区間$[-1,1]$を何個の離散値にするかを決める。ここでは`N=101`とする。
2. 区間$[-1,1]$を`N`個の離散値にした`ndarray`を定義する。
3. 2で作った`ndarray`の各要素値に対して`sin(2pi*x)`関数の値を求める。
4. 求めた値をmatplotlibでグラフにする。

```python
x = np.linspace(-1, 1, 101)  # N=100ではなく101にした理由を考えること。
y = np.sin(2 * np.pi * x)  # xの各要素に対してsin関数が適用される

fig = plt.figure()
ax = fig.add_subplot(111)
ax.plot(x, y)
# ax.plot(y) # yのみを引数にする場合も試して，x軸の表示がどう変わるか確認しなさい
```

Q. 以下のような関数`sinrnd()`を作りなさい。

概要: $x\in[0,3]$における$y=\sin 2\pi x$ の値に乱数を加えた値を要素とする`ndarray`を出力する。

- 引数
    - `nd`: 区間$[0,3]$を何個の離散値に分割するかを決める値(端点も含む)。デフォルト値は`101`。
- 処理
    1. 区間$[0,3]$を`num`個の離散値に分割した値を要素とする`ndarray`である変数`x`を生成する。
    2. `x`の各値に対する$y=\sin 2\pi x$の値を要素とする`ndarray`である変数`y`を生成する。
    3. `y`の各要素値に対して，平均0, 標準偏差1の正規分布乱数を加算する。
- 返り値
    - 上記処理により得た`x`と`y`を返り値とする。
- 注意
    - `x`の定義域に気をつけること

<!-- grade_id: cell-9a459e86c95b0c9d, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

作成した関数`sinrnd()`の返り値をグラフにすると，sin関数の形をとどめていない形になることが多い。次のセルを実行し，そのことを確認しなさい。

```python
x, y = sinrnd()
fig = plt.figure()
ax = fig.add_subplot(111)
ax.plot(x, y)
```

<!-- grade_id: E-sinrnd, points: 2, grade: true, solution: false -->
```python
# 評価用セル
```

上記のようにノイズの背後にある規則的な信号を取り出すには，複数回取得した信号の加算平均をとる。$n$個のデータの加算平均を行うと，S/N比はおよそ$\sqrt{n}$倍になる。

Q. `sinrnd()`で長さ`nd`の`ndarray`を指定した回数生成し，その各時刻(`x`の各値)における平均値と標準偏差を返す関数`mean_sinrnd()`を作りなさい。

**引数**
- `num`: `sinrnd()`を実行する回数。デフォルト値は`10`
- `nd` : `sinrnd()`に引数として渡すための引数。デフォルト値は`101`

**返り値**
1. `x` : `sinrnd(nd)`の返り値`x`
2. 平均値: `num`回`sinrnd(nd)`を実行して得られた返り値`y`の全てに対して，各時刻(`x`の各値)における平均値を要素とする`ndarray`
    - 例えば`num=2`の場合，`sinrnd(nd)`の1回目の返り値が`[0,0,1,1,...]`，2回目の返り値が`[4,4,5,5,]`なら，`mean_sinrnd()`の返り値は`[2,2,3,3,...]`になる。
3. 標準偏差: 上記の平均値の代わりに標準偏差を要素とする`ndarray`

**ヒント**
1. 0行`nd`列の空配列を生成する。その変数名は`dat`とする。**空配列の生成方法は以下の通り**。
```python
dat=np.empty((0,nd))
```
2. `sinrnd(nd)`の返り値を取得して`dat`の行に追加する。このとき，**返り値を1行`nd`列の2次元配列に変換**することを忘れない。これを`num`回実行。
3. 配列`dat`に対して平均と標準偏差を`np.mean()`と`np.std()`で計算。

<!-- grade_id: cell-ef7f990310cb621b, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-mean_sinrnd1, points: 1, grade: true, solution: false -->
```python
# 評価用セル
nd = 51
x, mean, std = mean_sinrnd(num=100, nd=nd)
eq_(len(x), nd)
eq_(len(mean), nd)
eq_(len(std), nd)
ans_x = np.linspace(0, 3, nd)
nt.assert_array_equal(x, ans_x)
```

<!-- grade_id: E-mean_sinrnd2, points: 1, grade: true, solution: false -->
```python
# 評価用セルその2
```

Q. 関数`mean_sinrnd()`の返り値を以下のようなグラフに示す関数`plot_sinrnd()`を作りなさい。以下のグラフは，`sinrnd()`を2000回呼び出して加算平均した場合の例。

**引数**
- `x`, `mean`, `std`: いずれも`mean_sinrnd()`の返り値

**返り値**

- `ax`: `mean`, `mean±std`を描画するための座標コンテナ


**注意**

- 各座標軸にタイトルをつけること
- 以下の図の凡例中の$\pm$は`$\pm$`で表示可能

**実行例**
```python
x, mean, std = mean_sinrnd(num=2000, nd=101)
plot_data(x, mean, std)

```

![](images/mean_sinrnd.png)

<!-- grade_id: cell-15e66f46230f3a2c, points: 0, grade: false, solution: true -->
```python
# 解答欄


def plot_data(x, mean, std):
    fig = plt.figure(figsize=(7, 4))
    ax = fig.add_subplot(111)
    # axに，mean, mean+SD, mean-SD の順に登録すること
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
x, mean, std = mean_sinrnd(num=2000, nd=101)
plot_data(x, mean, std)
```

<!-- grade_id: E-plot_mean_sinrnd, points: 2, grade: true, solution: false -->
```python
# 評価用セル
nd = 51
x, mean, std = mean_sinrnd(num=2000, nd=nd)
ax = plot_data(x, mean, std)

ans_y = np.sin(2 * np.pi * x)
pc = LinePlotChecker(ax)
pc.assert_y_data_allclose([ans_y, ans_y + 1, ans_y - 1], atol=0.13)
```

## 要素の総和

`L=np.arange(1,11)`の要素の総和は以下のように求めることができる。

```python
L = np.arange(1, 11)
sum = 0
for x in L:
    sum += x
print(sum)
```

`numpy.sum()`は，ndarry配列の各要素の総和を返す関数。これを使うと，上記のプログラムは以下のように簡潔になる。

```python
sum = np.sum(L)
print(sum)
```

**Q**. 以下のような関数`softmax(L)`を作りなさい。

- 引数:
    - `L`: `ndarray`配列
- 返値:
    - 引数`L`の要素を$l_0,l_2,\cdots,l_n$とおいた場合，$\displaystyle \frac{e^{l_0}}{\displaystyle\sum_{i=0}^ne^{l_i}}，\frac{e^{l_1}}{\displaystyle\sum_{i=0}^ne^{l_i}},\cdots,\frac{e^{l_n}}{\displaystyle\sum_{i=0}^ne^{l_i}}$を要素とする`ndarray`
    - 例) 引数が`L=[1,3,4]`の場合は，以下を計算した値が返り値。

        $\displaystyle[\frac{e^{1}}{S},\frac{e^{3}}{S},\frac{e^{4}}{S}]$, ただし，$S=e^{1}+e^{3}+e^{4}$

実行例
```python
x = np.array([0.0, 1.0])
print(softmax(x))
y = np.array([-1.0, 0.0, 1.0])
print(softmax(y))
```
実行結果
```python
[0.26894142 0.73105858]   # [1/(1+e), e/(1+e)]
[0.09003057 0.24472847 0.66524096] # [(1/e)/(1/e+1+e), 1/(1/e+1+e), e/(1/e+1+e)]
```

<!-- grade_id: cell-5c010430e99db7b9, points: 0, grade: false, solution: true -->
```python
# このセルに関数の定義を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしなさい。
```

<!-- grade_id: E-softmax, points: 2, grade: true, solution: false -->
```python
# 評価用セル
x = np.array([0.0, 1.0])
nt.assert_equal(softmax(L=x), np.array([1.0, np.e]) / (1.0 + np.e))
```

## softmax法
**Q.** 以下のような関数`softmax_rand(L,T)`を作りなさい。

- 引数:
    - `L`: リスト(要素は数値)
    - `T`: softmax法のパラメータ(デフォルト値は1)
- 返値:
    - 引数のリスト`L`の各要素値に対してsoftmax法を適用した確率で、そのインデックスをランダムに返す。
    - 例) 引数が`L=[1,3,4]`の場合は，ランダムに各要素のインデックスを返すが、インデックス`k`が選択される確率$P(k)$ ($k=0,1,2$)は次式の通り。
    
    $$\displaystyle P(k)=\frac{e^\frac{L(k)}{T}}{\displaystyle\sum_{i=0}^{N-1}e^\frac{L(i)}{T}}$$

実行例:
```python
import math as ma

softmax_rand(L=[1,1])       #返り値は0か1であり、その比率は1:1
softmax_rand(L=[0,1])       #返り値は0か1であり、その比率はexp(0):exp(1)=0:e
softmax_rand(L=[0,4],T=2)       #返り値は0か1であり、その比率はexp(0):exp(4/2)=0:exp(2)
softmax_rand(L=[0,ma.log(2),1]) #返り値は0か1か2であり、その比率は1:2:e
```
ヒント
1. 引数のリスト`L`を`ndarray`に変換する。
2. `L`を前問と同様にしてsoftmaxを用いた値に変換(要素値の合計が1になるように規格化される)。その結果で`L`をおき直す。
3. `[0,1)`の乱数を発生する。乱数値が`[0,L[0])`の範囲なら返り値は0，`[L[0],L[0]+L[1])`の範囲なら1， `[L[0]+L[1], L[0]+L[1]+L[2])`の範囲なら2,...と返り値を決定すれば良い。
    - `L`の要素数がいくつであっても対応できるようにすること
    - (独り言: 発生した乱数`rnd`に対して，(1) `L[0]`と比較, (2) `L[0]`より大きければ`rnd`から`L[0]`を引いて`L[1]`と比較...としていくと，簡潔なコードになる)

```python
# 乱数発生に利用するrandomモジュールの読み込み
import random as ra

# 対数計算に利用するmathモジュールの読み込み
import math as ma
```

<!-- grade_id: cell-cef0819b7706944b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

**Q2.** 前問で作った`softmax_rand()`のテストのため，以下のような関数`test()`を作りなさい。

- 概要
    - `num`回`softmax_rand()`を呼び出し，その返り値の分布を返り値とする
- 引数
    - `L`: リスト(要素は数値)
    - `num`: 確率発生回数(デフォルト値は適宜決めなさい)
    - `T`: softmax法のパラメータ(**デフォルト値は1**)    
- 処理
    - `num`回`softmax_rand(L,T)`を呼び出し，その返り値が引数`L`の各要素のインデックスにそれぞれ何回該当したかを記録する。
- 返り値
    - `num`回`softmax_rand(L,T)`を呼び出した時の返り値の確率分布
    - 型は`numpy.ndarray`
    - 例えば，`softmax_rand(L,T)`を10回呼び出した時の返り値が，`0`が2回，`1`が8回ならば返り値は`[0.2,0.8]`
    
ヒント

1. 引数のリスト`L`と同じ長さの`ndarray`を準備する。要素値は全て0とし，変数名を`count`とする。
2. 引数で指定された回数(`num`)だけ`softmax_rand(L,T)`を呼び出す。返り値が`0`なら`count[0]`の値を増加。
3. 最後に`count`の要素を，全要素の合計が1になるように規格化して返り値とする。

<!-- grade_id: cell-ef3c80ec5c89fcf9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Q. 関数`test()`を使って以下を確認しなさい。
- `softmax_rand(l=[1, 1], T=1)`の返り値は0か1であり、その比率はほぼ1:1である。
- `softmax_rand(l=[0, log2, 1], T=1)`の返り値は0か1か2であり、その比率はほぼ1:2:eである。

```python
# このセルで試してみよう
```

<!-- grade_id: cell-794b7bdf6a29341a, points: 1, grade: true, solution: false -->
```python
# 評価セル
n = 50000
## len(L)=2
r = test(L=[1, 1], num=n)
ans = np.array([0.5, 0.5])
nt.assert_array_almost_equal(r, ans, 0.05)

## len(L)=10
r = test(L=np.ones(10), T=2, num=n)
ans = np.full(10,0.1)
nt.assert_array_almost_equal(r, ans, 0.05)
```

<!-- grade_id: E-softmax_rand, points: 2, grade: true, solution: false -->
```python
# 評価用セルその2(その他いろいろなパタン検証)
```

```python

```