<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"   # 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a099bc"          # 大学アカウントのユーザ名
STUDENT_ID = "20-2202-000-0" # 省略なしで入力
```

```python
FAMILY_NAME = ""   # 姓をローマ字で入力
FIRST_NAME = ""   #  名をローマ字で入力
USER = ""          # 大学アカウントのユーザ名
STUDENT_ID = ""  # 12-2202-000-0 等と省略なしで入力
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">

# MDPGridworld

</div>

- レポート課題の準備のための問題です。レポートはこのシートで準備したクラスや関数を用いて検討・考察した内容をまとめていただきます。

## `MDPGridworld-v1`:

- [openAI Gym](https://gym.openai.com/)準拠のモジュール。迷路等をデザインして強化学習環境として利用できる。
- ソースコードは[こちら](https://github.com/jnishii/gridworld-gym)。([IRLL/reinforcement_learning_class](https://github.com/IRLL/reinforcement_learning_class/tree/master/gym/envs)のものを改変したもの)

MDPGridworldのモジュールの読み込みは次セルを実行

```python
import gym

env = gym.make("MDPGridworld-v1")
```

## 状態と行動

### 状態(state)
エージェントは，自分がどの部屋にいるか(部屋ID)を状態(state)の情報として知ることができる。
各部屋の配置および，状態(部屋ID)は以下の通り。

```python
env.show_info()
```

### 行動(action)

選択できる行動は以下に示す0から3の4通り。もし壁に向かって移動しようとしたら，その場に留まる。
```python
0 - Northに移動
1 - Southに移動
2 - Westに移動
3 - Eastに移動
```

### エピソード(episode)とエージェントの目的

- **1エピソード**: スタート(S)から，ゴール(G)もしくはFの部屋に到達するまで
- **エージェントの目的**: 1エピソード間にできるだけ多くの報酬を得ること

## openAI gymのインタンス変数・関数
### 状態に関する情報を格納するインスタンス変数
- `gym.observation_space` : 状態に関する情報
- `gym.observation_space.n`: 状態数

次セルでは各インスタンス変数の値を確認している。

```python
print(env.observation_space)
print(env.observation_space.n)
```

実行結果より，状態が離散的な64個の値(0から63まで)をとることがわかる。

### 行動に関するインスタンス変数
- `gym.action_space`: 行動に関する情報
- `gym.action_space.n`: 行動数
- `gym.action_space.sample()`: エージェントがとりうる行動を一つランダムに選択して返す

各インスタンス変数の値は以下の通り。

```python
print(env.action_space)
print(env.action_space.n)
print(env.action_space.sample())
```

実行結果より，とりうる行動は離散的な4個の値(0から3まで)で表現されていることがわかる。また，`env.action_space.sample()`の返り値は，この４個の値からランダムに選ばれている。

**Q.** 上述の`MDPGridWorld-v1`環境には，状態と行動の組み合わせは全部何通りあるか? その値を変数`ans`に代入しなさい。

<!-- grade_id: cell-cc03c48248013b48, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-gridworld-states, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

### その他の情報

`gym.spec.max_episode_steps`
    - 1エピソード内の最大ステップ数
    - ゴールに到達するか，もしくはこのステップ数に達したところでエピソードは終了
**Q.** 上記のインスタンス変数の値を確認しなさい。

```python
print(env.spec.max_episode_steps)
```

### 環境クラスのインスタンス・メソッド

- `env.reset()`: 初期化。エージェントの位置はSになる。
- `env.step(action)`: 
    - 引数:
        - `action` : エージェントのとる行動
    - 返り値:
        - `next_state` : actionによる状態遷移後の状態
        - `reward` : actionにより得た報酬
        - `terminal` : エピソード終了時にTrue, それ以外はFalse
        - `info` : 行動actionによって，状態sから次の状態next_stateに遷移する確率(この演習では使わない)

次セルで`env.reset()`を実行し，適当な`action`を用いて`env.step()`による状態遷移を行い，`action`に応じた適切な状態，報酬値，終了状態が得られることを確認しなさい。

```python
# 以下のactionを変更し，actionに応じたいろいろな返り値が得られていることを確認しなさい。
env.reset()
action = 0
next_state, reward, terminal, info = env.step(action)
# terminal is a return value of the environment's step method
print(next_state, reward, terminal, info)
```

<div class="alert alert-success">

# MDPGridWorldを動き回る

</div>

## ランダムアクション
前回の課題(rooms課題)で作ったクラス`Fool`および，`play_env()`の修正版を用いて`MDPGridWorld-v1`の環境を動き回るプログラムを作る。


**Q(1)** 以前の課題(rooms課題)で作ったクラス`Fool`の定義を以下に入力しなさい。

```python
# このセルでは，実行に必要なモジュールの読み込みを行う。
# ここに書かれているもの以外に必要なものがあれば，適宜追加すること。
import numpy as np
import random as ra
import matplotlib.pyplot as plt
```

<!-- grade_id: cell-08d712ac2ed07e1c, points: 0, grade: false, solution: true -->
```python
# class Foolの定義をこのセルに入力する
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-gridworld-Fool, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

**Q(2)** 以前の課題(rooms課題シート)で作った`play_env()`を以下のよう修正しなさい。

**変更点**
1. 引数`max_step`(1episode内の最大ステップ数)は廃止。`play_env()`内での処理には，代わりに`env.spec.max_episode_steps` の値を用いること。
2. 引数`s0`も廃止
3. 関数`env.step(action)`の返り値はこれまで`state,reward,terminal`の3つだったが，openAI gymでは4つめに`info`が加わる。ただ，この引数の内容を使うことはほとんど無いので，以下のように返り値を受け取って，4つ目の引数は無視する。
```python
s,r,terminal,_=env.step(action) # 返り値の第4要素は'_'という名前の変数で受け取っている。
```
`play_env()`を呼び出して，ランダム選択エージェント`Fool`が`FronzenLake-v1`の環境を動き回る実行例は，解答欄の次のセルにある。

<!-- grade_id: cell-62f3023136ba8250, points: 0, grade: false, solution: true -->
```python
def play_env(env, agent, learning=True, max_episode=30, render_step=20, verbose=False):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# `play_env()`を呼び出して，ランダム選択エージェントFoolが`MDPGridworld-v1`の環境を動き回る実行例。
# 引数をいろいろと変えて，動作検証をしなさい。

env = gym.make("MDPGridworld-v1")
action_list = list(range(env.action_space.n))
step, score = play_env(
    env,
    agent=Fool(action_list),
    learning=False,
    max_episode=1,
    render_step=1,
    verbose=True,
)
```

<!-- grade_id: E-gridworld-play_env, points: 1, grade: true, solution: false -->
```python
# 評価用セル(Foolを用いた動作確認。demo_env()と同様の動作確認)
env = gym.make("MDPGridworld-v1")
action_list = list(range(env.action_space.n))


def test_mdpgrid(agent, verbose=True):
    return play_env(
        env=env,
        agent=agent,
        learning=False,
        max_episode=1,
        render_step=1,
        verbose=verbose,
    )


result = [test_mdpgrid(Fool(action_list, 0), verbose=False) for i in range(2)]
eq_(result, ([([200], [-600]), ([200], [-600])]))
result = [test_mdpgrid(Fool(action_list, 1), verbose=False) for i in range(2)]
eq_(result, ([([200], [-600]), ([200], [-600])]))
result = test_mdpgrid(Fool(action_list, 2), verbose=False)
eq_(result, ([200], [-600]))
result = test_mdpgrid(Fool(action_list, 3), verbose=False)
eq_(result, ([200], [-600]))
result = [test_mdpgrid(Fool(action_list, "random"), verbose=False) for i in range(10)]
result2 = [test_mdpgrid(Fool(action_list, "random"), verbose=False) for i in range(10)]
assert_not_equal(result, result2)
```

以前の課題(rooms課題)で作った関数`plot_history()`を使えば，エピソードごとのステップ数や総得点の履歴のグラフ表示もできる。

**Q.** 次のセルにclass課題で作った`plot_history()`を入力して動作確認を行い，必要に応じて修正をしなさい。動作確認のための実行例は解答欄の次のセルにある。

注意)
- class課題で作った関数`plot_history()`で定義した`ax1`,`ax2`の変数名は変更しないこと。
- 返り値も`ax1, ax2`のまま変更しないこと。

<!-- grade_id: cell-451de4f61c742e83, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt


# YOUR CODE HERE
raise NotImplementedError()
```

```python
# 実行例
# 引数をいろいろと変えて，動作検証をしなさい。
env = gym.make("MDPGridworld-v1")
step, score = play_env(
    env,
    agent=Fool(action_list),
    learning=False,  # Foolは学習しない
    max_episode=100,
    verbose=False,
)
plot_history(step, score)
```

<!-- grade_id: E-gridworld-plot_history, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## グラフのファイル保存

前問のブログラムの場合，`plt.tight_layout()`の直後に以下の命令を追加すると**グラフをファイル保存**できる。
```python
plt.savefig('figure.png')
```
保存されたファイルは，課題シートの一覧画面で確認やダウンロードをできる。

## Q学習
**Q.** 以前の課題(rooms)で作ったクラス`Qlearner`と，このシートで定義済みの関数`play_env()`を使えば，MDPGridWorld環境に対する学習もできるはずである。次のセルに`Qlearner`の定義を入力して動作確認を行い，必要に応じて修正しなさい。

実行例は解答欄の次のセルにある。

<!-- grade_id: cell-3e4bacb41b3355e1, points: 0, grade: false, solution: true -->
```python
# 以下にclass QLearnerの定義を書きなさい。(必要ならindex_list()等, Qlearnerの実行に必要な関数も書きなさい)
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

<!-- grade_id: E-gridworld-Qlearner, points: 1, grade: true, solution: false -->
```python
# 評価用セル
```

## 学習の継続
Q学習エージェント`agent`の学習結果(Q値)は，関数`play_env()`を終了しても変数`agent`に保持される。したがって，上のセルの実行後に，以下のように`play_env()`を再度実行すれば，上記の学習の続きを行うことができる。
```python
step,score=play_env(env, agent=agent, max_episode=100, verbose=False)
plot_history(step,score)
```
**学習結果(Q値)を未学習な状態に戻したいとき**は，`agent`を再定義するか，`agent.reset()`を実行する。

**Q.** 実際に`play_env()`を何度か実行してみて，未学習の状態よりも，パフォーマンスが良くなっていくことを確認しなさい。

```python
step, score = play_env(env, agent=agent, max_episode=100, verbose=False)
plot_history(step, score)
```

**Q.** Q学習エージェントは学習により，どのようなQ値を獲得しただろうか。可視化を試みよ。例えば，Q値を格納している辞書`agent.Q`の内容を見ると，獲得したQ値の一覧を確認できる。以下は，可視化のサンプルプログラムだが，このサンプルプログラムをレポート等に利用したい場合には，座標軸の値やタイトルを加え，また，表示内容に関する十分な解説を加えて使うこと。プログラムの内容を理解しないまま，図だけ見て意味を判断して使うと，勘違いをすることがあるので要注意。

```python
def show_Q(agent, n_col=8):
    Qarray = []
    Qrow = []
    nCol = n_col
    print("+--------+")
    print("|", end="")
    for s in range(env.observation_space.n):
        q = [agent.get_q(s, a) for a in range(env.action_space.n)]
        Qrow.append(max(q))
        argmaxQ = q.index(max(q))  # Q値が最大になるaのうち, indexが最小のもの
        print("{}".format(["^", "v", "<", ">"][argmaxQ]), end="")
        if (s + 1) % nCol == 0:
            print("|")
            if s != env.observation_space.n - 1:
                print("|", end="")
            Qarray.append(Qrow)
            Qrow = []
    print("+--------+")

    fig = plt.figure(figsize=(6, 5))  # グラフを生成
    ax = fig.subplots()  # 座標軸を1つ生成してaxに代入
    ax.invert_yaxis()    # y軸を上下ひっくり返したいとき

    tick = range(0, 8, 1)
    ax.set_xticks(tick)
    ax.set_yticks(tick)

    map = ax.pcolor(Qarray, cmap="bwr")  # 数値データからカラーデータを作成 cmapの値により配色変化
    fig.colorbar(map, ax=ax)  # カラーデータと軸データを使ってグラフデータ生成

    plt.show()  # 表示(生成されているグラフデータを画面表示)


show_Q(agent)
```

Q. `class Qlearner`の継承クラスとして，`class SARSAlearner`を作り，学習のためのインスタンス関数`learning()`の内容をSARSA学習にしなさい。

<!-- grade_id: cell-d04558d0637745a7, points: 0, grade: false, solution: true -->
```python
# このセルにクラス定義を入力する
class SARSAlearner(Qlearner):
    def learning(self, s1, a1, r, s2, terminal=False):
        """
        Q(s1, a1) <= Q(s1, a1) + alpha * (r + gamma*Q(s2,a') - Q(s1,a1)) # 通常時
        Q(s1, a1) <= Q(s1, a1) + alpha * (r - Q(s1,a1)) # エピソード終了時
        a'は，状態s'において行動選択ポリシーに従い選択される行動(get_action()で取得)
        """
        # YOUR CODE HERE
        raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
# SARSAlearner + play_env() + MGPGridWorld の学習プログラム

agent = SARSAlearner(action_list)  # 学習エージェント生成
step, score = play_env(env=env, agent=agent, max_episode=100, verbose=False)
plot_history(step, score)
```

<!-- grade_id: E-gridworld-SARSAlearner, points: 1, grade: true, solution: false -->
```python
# 評価用セル
class SARSAlearnerTest(SARSAlearner):
    def set_action(self, a=0):
        self.def_action = a

    def get_action(self, s):
        return self.def_action


env = gym.make("MDPGridworld-v1")
action_list = list(range(env.action_space.n))  # action_list取得
agent = SARSAlearnerTest(action_list)  # 学習エージェント生成

# test1
agent.set_action(a=0)
step, score = play_env(env=env, agent=agent, max_episode=5, verbose=False)
ans_Q = {
    (56, 0): -2.882376096,
    (48, 0): -2.882376096,
    (40, 0): -2.882376096,
    (32, 0): -2.9070494918446594,
    (24, 0): -3.3473247916120967,
    (16, 0): -6.290059022035771,
    (8, 0): -15.00201805496299,
    (0, 0): -24.509190126958917,
}
eq_(agent.Q, ans_Q)
# test2
```

<div class="alert alert-danger">

# ここから下には，プログラムを書かないでください

</div>    
    
プログラムのテストは，これまでに提出済みのシートに追加セルを作成して行ってください。